Superoxide Dismutases maintain niche homeostasis in stem cell populations

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

In Figure 1, it is very difficult to identify where CySCs end and GSCs begin without using a cell surface marker for these different cell types. In addition, the methods for quantifying the mitochondrial distribution in GSCs vs. CySCs are very much unclear and appear to rely on colocalization with molecular markers that are not in the same cellular compartment (Tj-nuclear vs. Vasa-perinuclear and cytoplasmic) the reader has no way to determine the validity of the mitochondrial distribution. Similarly, the labelling with gstD1-GFP is also very much unclear - I see little to no GFP signal in either GSCs or CySCs in panels 1GK. Lastly, while the expression o SOD in CySCs does increase the gstD1-GFP signal in CySCs, the effects on GSCs claimed by the authors are not apparent.

We appreciate the reviewer’s detailed feedback on Figure 1 and the concerns raised regarding identifying CySCs and GSCs, as well as the methods used for quantifying mitochondrial distribution and gstD1-GFP labeling. Below, we address each point and describe the revisions made to improve clarity and rigor

Distinguishing CySCs and GSCs and Mitochondrial Distribution in GSCs vs. CySCs in Figure1

We acknowledge the difficulty in distinguishing CySCs from GSCs without the use of additional cell surface markers. To improve clarity, we have now included a membrane marker discslarge (Dlg) in our revised Figure 1 and S1 to delineate cell boundaries more clearly. Additionally, we provide higher-magnification images to indicate the mitochondria in CySCs and GSCs. We also agree that ing on mitochondrial distribution might be far-fetched. In the revised manuscript, we have limited our analysis to mitochondrial shape, which was found to be different in GSC and CySC (Fig. 1, D, F, G, and S1B). We have clarified our quantification methods in the revised Methods section, providing details on the image processing and analysis pipeline used to assess mitochondrial distribution.

Clarity of gstD1-GFP Labelling:

We recognize the reviewer’s concern regarding the weak GFP signal in these panels. To improve visualization, we have included fresh set of images by optimizing the contrast and presenting additional monochrome images with higher exposure settings to better illustrate gstD1-GFP expression (Figure 1L,1Q, and S1C’’’-D’’’). Additionally, we have demarcated the cell boundaries using Dlg along with individual labelling of Vasa+ and Tj+ cells. Due to technical difficulty associated with acquisition of images, we could not co-stain Vasa, Tj and Dlg together. Therefore, quantified the gstD-GFP intensity separately for GSCs and CySCs under similar acquisition conditions (Figure 1R).

Effects of SOD depletion on GSCs:

While our initial analysis suggested changes in gstD1-GFP expression in GSCs upon Sod1 depletion in CySCs, we acknowledge that the effects may not be as apparent in the provided images. In response, we have expanded our quantification, included a statistical analysis of gstD1-GFP intensity specifically in GSCs and CySCs (Figure 1S), and added more representative images in the revised figure panels (Figure S1C-D’’’) to support our claims.

In Figure 2, while the cell composition of the niche region does appear to be different from controls when SOD1 is knocked down in the CySCs, at least in the example images shown in Figures 2A and B, how cell type is quantified in figures 2E-G is very much unclear in the figure and methods. Are these counts of cells contacting the niche? If so, how was that defined? Or were additional regions away from the niche also counted and, if so, how were these regions defined?

Thank you for your regarding the quantification of cell types in Figures 2E-G. We counted all cells that were Tj-positive and Zfh1-positive in individual testis, while for GSCs, only those in direct contact with the hub were included. This clarification has been incorporated into the revised figure legend and methods (line no.400-407). We have now provided a clearer description in the text to improve transparency in our analysis.

In Figure 3, it is quite interesting that there is an increase in Eya+, differentiating cyst cells in SOD1 knockdown animals, and that these Eya+ cells appear closer to the niche than in controls. However, this seems at odds with the proliferation data presented in Figure 2, since Eya+ somatic cells do not normally divide at all. Are they suggesting that now differentiating cyst cells are proliferative? In addition, it is important for them to show example images of the changes in Socs36E and ptp61F expression.

Thank you for your insightful observations. We acknowledge the apparent contradiction and appreciate the opportunity to clarify our interpretation.

Regarding the increase in Eya+ differentiating cyst cells in Sod1RNAi individuals and their proximity to the niche, we do not suggest that these differentiating cells are proliferative. Instead, we propose that the knockdown of Sod1 may alter the timing or regulation of cyst cell differentiation, leading to an accumulation of Eya+ cells near the niche. To clarify this point, we have revised the manuscript (line no. 186-189) to emphasize that our proliferation data specifically refers to early-stage somatic cells, not Eya+ differentiating cyst cells.

We also appreciate the reviewer's request for example images illustrating the changes in Socs36E and Ptp61F expression. We could not access the antibodies specific to Socs36E and Ptp61F. Hence, we had to rely on the measurements were obtained using real-time PCR from the tip region of testis. We have clarified the same in the figure legends (line 700).

Overall, the various changes in signaling are quite puzzling-while Jak/Stat signaling from the niche is reduced, hh signaling appears to be increased. Similarly, while the authors conclude that premature differentiation occurs close to the niche, EGF signaling, which occurs from germ cells to cyst cells during differentiation, is decreased. Many times these, changes are contradictory, and the authors do not provide a suitable explanation to resolve these contradictions.

We appreciate the reviewer’s thoughtful feedback on the signaling changes described in our study. We acknowledge that the observed alterations in Jak/Stat, Hedgehog (Hh), and EGF signaling may appear contradictory at first glance. However, our data suggest that these changes reflect a complex interplay between different signaling pathways that regulate cyst cell behavior in response to specific genetic perturbation.

Regarding Jak/Stat and Hh signaling, while Jak/Stat activity is reduced in the niche, the increase in Hh signaling may reflect a compensatory mechanism or a context-dependent response of cyst cells to reduced Jak/Stat input. Prior studies have suggested that Hh signaling can function in parallel and independently of Jak/Stat signaling (PMID: 23175633) and our findings align with this possibility.

The reduction in EGFR signaling in this context appears contradictory to existing literature. One possible explanation is that, the altered GSC -CySC balance and loss of contact in Tj>Sod1i testes, leads to insufficient ligand response, thereby failing to activate EGFR signaling. (line no.222-224, 313-318).

Reviewer #2 (Public review):

We sincerely appreciate the reviewer’s detailed feedback, which has helped refine our manuscript. In this study we have focussed on the role of ROS generated due to manipulation of Sod1 in the interplay between GSC and CySCs. In this regard, we have conducted additional experiments and incorporated quantitative data into the revised manuscript. Additionally, we have refined the text and provided further context to enhance the clarity. Key revisions include:

(1) Clarification of Quantification Methods – We have refined intensity measurements by incorporating a membrane marker (Dlg) to better delineate cell boundaries and have normalized Ptc and Ci expression per cell to improve clarity.

(2) Cell-Specific ROS Measurement – We separately measured ROS in germ cells and cyst cells and performed independent Sod1 depletion in GSCs to determine its direct effects.

(3) Mitochondrial Analysis – We revised our approach, focusing on mitochondrial shape rather than asymmetric distribution, and removed overreaching claims.

(4) Proliferation Analysis – We reanalyzed FUCCI data by normalizing to total cell count, supporting the conclusion that increased proliferation, rather than differentiation delay, underlies the observed phenotype.

(5) E-Cad Quantification – We specifically analyzed E-Cad levels at the GSC-hub interface to strengthen conclusions on GSC attachment.

(6) JAK/STAT Signaling – While we could not obtain a STAT92E antibody, we clarified the spatial limitations of our current analysis and revised the text accordingly.

(7) Rescue Experiments and Gal4 Titration Control – We performed additional control experiments to confirm that observed effects are not due to Gal4 dilution.

(8) Image Quality and Terminology Corrections – We enhanced figure resolution, corrected terminology (e.g., "cystic" to "cyst"), and revised ambiguous phrasing for clarity and accuracy.

As suggested, we have also changed the manuscript title to better align with our results:

Previous Manuscript Title: Non-autonomous cell redox-pairs dictate niche homeostasis in multi-lineage stem populations

Updated Manuscript Title: Superoxide Dismutases maintain niche homeostasis in stem cell populations

Specific responses to the reviewer’s:

While the decrease in pERK in CySCs is clear from the image and matched in the quantification, the increase in cyst cells is not apparent from the fire LUT used. The change in fluorescence intensity therefore may be that more cells have active ERK, rather than an increase per cell (similar arguments apply to the quantifications for p4E-BP or Ptc). Therefore, it is hard to know whether Sod1 knockdownresults in increased or decreased signaling in individual cells.

Thank you for your insightful . To clarify, in the Fire LUT images, only pERK intensity is shown, not the cyst cell number. In our context, while there are more cells, the overall pERK intensity is lower, eliminating any ambiguity about whether the change is occurring per cell or due to an increased number of circulating cells. Moreover, for Ptc and Ci levels, we have normalized Ptc and Ci expression intensity per cell to enhance clarity and ensure an accurate interpretation of signaling changes.

There are several places in which the authors could strengthen their manuscript by explaining the methods more clearly. For example, it is unclear how the intensity graphs in Figure 1Q are obtained. The curves appear smoothed and therefore unlikely to be from individual samples, but this is not clearly explained. However, this quantification method is clearly not helpful, as it shows the overlap between somatic and germline markers, suggesting it cannot accurately distinguish between the two cell types. Additionally, using a nuclear marker (Tj) for the cyst cells and cytoplasmic marker (Vasa) for the germ cells risks being misleading, as one would not expect much overlap between cytoplasmic gstD1-GFP and nuclear Tj. Also related to the methods, it is unclear how Vasa+ cells at the hub were counted. The methods suggest this was from a single plane, but this runs the risk of being arbitrary since GSCs can be distributed around the hub in 3D. (As a note, the label on the graph "Vasa+ cells" is misleading, as there are many more cells that are Vasa-positive than the ones counted.)

We appreciate the reviewer’s careful evaluation of our manuscript and their insightful suggestions for improving the clarity of our methods. Below, we address each concern raised and describe the revisions made accordingly.

Clarification of Intensity Graphs in Figure 1Q

We have removed this graph, as we recognize that the markers previously used were not appropriate for distinguishing the different cell types. To address this concern, we have revised the text and now included a membrane marker discs-large (Dlg) in our revised Figure 1 and S1 to more clearly delineate cell boundaries. Due to technical difficulty associated with acquisition of images, we could not co-stain Vasa, Tj and Dlg together. Therefore, quantified the gstD-GFP intensity separately for GSCs and CySCs under similar acquisition conditions (Figure 1R).

Counting of Vasa+ Cells at the Hub

We appreciate the reviewer’s concern regarding our method for counting Vasa+ cells. In our original analysis, we included GSCs as the Vasa-positive cells that were in direct contact with the hub. To account for the three-dimensional arrangement of GSCs, we used the Cell counter plugin of Fiji and performed counting across different focal planes to ensure all hub-associated cells were considered. For better clarity on cell distribution around the hub, we have presented a single focal place image sliced through mid of the hub zone. To enhance transparency, we have now provided a more detailed explanation of our counting approach in the Methods section (line no 400- 403).

We agree that the label "Vasa+ cells" may be misleading, as many cells express Vasa beyond the specific subset being counted. To address this, we have changed the label to " GSCs" to reflect the subset analyzed more accurately.

The crucial experiment for this manuscript is presented in Figures 1 G-S, arguing that Sod1 knockdown with Tj-Gal4 increases gstD1-GFP expression in germ cells. This needs strengthening as the current quantifications are not convincing and appear to show an overlap between Tj (a nuclear cyst cell marker) and Vasa (a cytoplasmic germ cell marker). Labeling cell outlines would help, or alternatively, labeling different cell types genetically can be used to determine whether the expression is increased specifically within that cell type. Similarly, the measurement of ROS shown in the supplemental data should be conducted in a cell-specific manner. To clearly make the case that Sod1 knockdown in cyst cells is impacting ROS in the germline, it would be important to manipulate germ cell ROS independently. Without this, it will be difficult to prove that any effects observed are a result of increased ROS in the germline rather than indirect effects on the germline of altered cyst cell behaviour.

We appreciate the reviewer’s insightful feedback regarding the specificity of Sod1 knockdown effects in germ cells and the need for clearer quantification in Figures 1G–S. Below, we address each concern and outline the modifications made:

Clarification of Cell Type-Specific Expression:

We acknowledge the overlap observed between Tj (nuclear cyst cell marker) and Vasa (cytoplasmic germ cell marker) in the presented images. To strengthen our claim that gstD1GFP expression increases specifically in germ cells upon Sod1 knockdown, we have now labelled cell outlines using membrane marker discs-large (Dlg) to better distinguish cell boundaries, along with individual labelling of Vasa+ and Tj+ cells. Due to technical difficulty associated with acquisition of images, we could not co-stain Vasa, Tj and Dlg together.

Cell-Specific Measurement of ROS:

We agree that a cell-type-specific ROS measurement is critical to establishing a direct effect on germ cells. To address this, we have now performed ROS measurements separately in germ cells and cyst cells under similar acquisition conditions. These data are now included in the revised (Figure 1R). Similarly, upon CySC-specific Sod1 depletion, we performed measurement of gstD1-GFP intensity which was found to be enhanced in GSCs, along with expected increase in CySCs (Fig 1S). We have independently manipulated ROS levels in GSCs (Nos Gal4> Sod1i) and observed that elevated ROS negatively impacts GSCs, leading to a reduction in their number, while having an insignificant effect on adjacent CySCs.(Fig S2 E, F).

Quantifications of mitochondrial localization in Figure 1 should include some adequate statistical method to evaluate whether the distribution is random or oriented towards the GSC/CySC interface. From the image provided (Figure 1B), it would appear that there are two clusters of mitochondria, on either side of a CySC nucleus, one cluster towards a GSC and one cluster away. Therefore evaluating bias would be important. Additional experiments will be necessary to support the statement that "Redox state of GSC is maintained by asymmetric distribution of CySC mitochondria". This would require manipulating mitochondrial distribution in CySCs.

We appreciate the reviewer’s suggestion regarding the quantification of mitochondrial localization. We agree that ing on mitochondrial distribution might be far-fetched. In revised manuscript, we have demarcated the cell boundary and limited our analysis to mitochondrial shape which was found to be different in GSC and CySC (Fig. 1, D, F, G and S1B). Mitochondrial shape was quantified based on the mitochondrial area and circularity (Figure 1F and G). To prevent any misinterpretation, we have removed the statement, "Redox state of GSC is maintained by asymmetric distribution of CySC mitochondria."

One point raised by the authors is that the increase of somatic cell numbers is driven by accelerated proliferation, based on an increased number of cells in various stages of the cell cycle as assessed by the FUCCI reporter. However, there are more somatic cells in this genetic background, so it could be argued that the observed increase in different phases of the cell cycle is due to an increased number of cells. In order to argue for an increased proliferation rate, the number of cells in each phase should be divided by the total number of cells, expecting to see an increase in S and G2/M phases along with a decrease in G1. Otherwise, the simplest explanation is a block or delay in differentiation, meaning that more cells remain in the cell cycle.

We appreciate the regarding the interpretation of our FUCCI reporter data. We acknowledge that the observed increase in the number of cells in various phases of the cell cycle could be influenced by the overall higher number of somatic cells in this genetic background.

To address this concern, we have now re-analyzed our FUCCI data by normalizing the number of cells in each phase to the total number of cells and we did not observe a significant shift in the proportion of cells in S and G2/M phases relative to G1. This suggests presence of more proliferative cells, that is less cells in Go phase, rather than alterations in the timing of cell cycle progression stages. We are not sure about a block in differentiation because we see an enhanced accumulation of Eya+ cells near the niche. We have also supported our FUCCI data with pH3 staining where we have found more pH3+ spots under SOD1 depleted background. We have revised our manuscript accordingly (Figure 2I, K and S2U) to reflect this interpretation and appreciate the constructive feedback.

In Figure 3, the authors claim that knockdown of Sod1 in the soma decreases the attachment of GSCs to the hub-based on lower E-Cad levels compared to controls. Previous work has shown that in GSCs, E-Cad localizes to the Hub-GSC interface (PMID: 20622868). Therefore, the authors should quantify E-Cad staining at the interphase between the germ cells and the niche.

We appreciate the reviewer’s . As suggested, we have now quantified ECad staining specifically at the interface between the germ cells and the niche. Our analysis confirms that E-Cad levels are significantly reduced at this interphase upon Sod1 knockdown in the soma compared to controls, supporting our conclusion that Sod1 depletion affects GSC attachment to the hub as well as the whole niche. The revised Figure 3M now includes these quantifications, and we have updated the figure legend and results section accordingly.

The authors show decreased expression of the JAK/STAT targets socs36E and ptp61F, arguing that this could be a reason for decreased GSC adhesion to the hub. However, these data were obtained from whole testes and lacked spatial resolution, whereas a STAT92E staining in control and tj>Sod1 RNAi testes could easily prove this point. Indeed, previous work has shown that socs36E is expressed in the CySCs, not GSCs (PMID: 19797664), suggesting that any decrease in JAK/STAT may be autonomous to the CySCs.

We appreciate the reviewer’s observation regarding the spatial resolution of our JAK/STAT target expression analysis. To improve accuracy, we have attempted to collect only the tip of the testes while excluding the rest; however, we acknowledge that this approach may still obscure cell-specific changes. We had attempted to procure the STAT92E antibody but, despite multiple inquiries, we did not receive a positive response. While we agree that STAT92E staining would have strengthen our findings, we are currently unable to perform this experiment. Nevertheless, our observations align with prior work indicating that socs36E is predominantly expressed in CySCs (PMID: 19797664). We have revised the manuscript text accordingly to clarify this limitation.

Additional considerations should be taken regarding the rescue experiments where PI3KDN and Hh RNAi are expressed in a Tj>Sod1 RNAi background. To rule out that any rescue can be attributed to titration of the Gal4 protein when an additional UAS sequence is present, a titration control would be useful. These pathways are not described accurately since Insulin signaling is necessary for the differentiation of somatic cells (not maintenance as written in the text), and its inhibition has been shown to increase the number of undifferentiated somatic cells (PMID:27633989). As far as Hh is concerned, the expression of this molecule is restricted to the niche. It would be important to establish whether the expression is altered in this case, especially as the authors rescue the Sod1 knockdown by also knocking down Hh. One possibility that the authors need to rule out is that some of the effects they observe are due to the knockdown of Sod1 (and/or Hh) in the hub as Tj-Gal4 is expressed in the hub as well as the CySCs (PMID:27546574).

We appreciate the reviewer’s insightful s and suggestions. Below, we address each concern and describe the steps we have taken to incorporate the necessary modifications in our revised manuscript.

Titration Control for Rescue Experiments

We acknowledge the reviewer’s concern regarding potential Gal4 titration effects when introducing additional UAS constructs. To address this, we conducted a control experiment quantifying SOD1 levels in control, Tj > Sod1 RNAi, and Tj > Sod1 RNAi, UAS hhRNAi backgrounds using real-time PCR (Figure S4 M). The Sod1 levels in single and double UAS copy conditions were comparable, indicating that Gal4 titration does not significantly affect the results.

Clarification of Insulin Signaling Role

We appreciate the reviewer’s insight regarding the involvement of insulin signaling in this context. Initially, we included data on PI3K/TOR as we found it intriguing. However, as the data didn’t add much to the overall observations, we have removed them to ensure clarity and prevent any potential confusion.

Hh Expression and Niche Consideration

We recognize the importance of evaluating whether Hedgehog (Hh) expression is altered in the Sod1 RNAi background. We have already quantified hh in qRT-PCR (Figure S4C).

Potential Effects of Sod1 and Hh Knockdown in the Hub

We acknowledge the concern that Tj-Gal4 is expressed in both the hub and CySCs, potentially affecting hub function upon Sod1 and Hh knockdown. To address this, we have included additional data using the CySC-specific driver C-587 Gal4 to distinguish CySC-intrinsic effects from potential hub contributions. Our results show that while the phenotypic changes are consistent across both drivers, the effects are significantly stronger with Tj-Gal4, suggesting a role of the hub in this process. These findings have been incorporated into the revised manuscript (Fig S1G-H, M-N).

In general, the GSCs (and other aspects) are difficult to see in the images; enlargements or higher-resolution images should be provided. Additionally, the manuscript contains several mistakes or inaccuracies (examples include referring to ROS having "evolved" in the abstract when it is cells that have evolved to use ROS, or the references to "cystic" cells when they are usually referred to as "cyst" cells, or that "CySCs also repress GSC differentiation by suppressing transcription of bag-of-marbles" when CySCs produce BMPs that lead to suppression of bam expression in the germline). These would need editing for both clarity and accuracy.

We appreciate the reviewer’s insightful feedback and have made the necessary revisions to address the concerns raised.

Image Clarity and Resolution:

We have provided higher-resolution images in some of the revised images for better understanding. The revised figures now offer better clarity for key observations.

Clarification of Terminology and Accuracy:

The phrase regarding ROS in the abstract has been revised to reflect that cells have evolved to utilize ROS, rather than ROS itself evolving (line no. 27).

References to "cystic" cells have been corrected to "cyst" cells for consistency with standard terminology.

The statement about CySCs repressing GSC differentiation has been revised for accuracy, clarifying that CySCs produce BMPs, which lead to the suppression of bam expression in the germline (line no. 84).

We have carefully reviewed the manuscript for any additional inaccuracies or ambiguities to ensure clarity and precision. We appreciate the reviewer’s constructive s, which have helped improve the manuscript.

Reviewer #3 (Public review):

In response to Reviewer 3’s comments, we would like to highlight the point that in the present study we have focussed on the interplay between CySC and GSC and have accordingly conducted our experiments. We did observe some changes in the hub and do not rule out the effect of hub cells in exacerbating some of our phenotypes. We have included additional controls to highlight the effect of CySC ROS. These points have been appropriately discussed in the manuscript. Key revisions include:

(1) Data Clarity & Visualization: To improve mitochondrial lineage association, we incorporated a membrane marker (Dlg) in Figure 1, enhancing the distinction between CySCs and GSCs. Additionally, we refined gstD-GFP quantifications in individual cell types and provided high-resolution images.

(2) ROS Transfer & Measurement: We revised our discussion to acknowledge indirect ROS transfer mechanisms and added separate ROS quantifications in GSCs and CySCs, confirming higher ROS levels in CySCs (Figure 1R).

(3) Tj-Gal4 Specificity & Niche Characterization: Recognizing Tj-Gal4 expression in hub cells, we included C587-Gal4 as a CySC-specific driver, demonstrating that hub cells contribute partially to the phenotype (Figure S1G,H,M,N).

(4) Signaling Pathway Validation: We optimized dpERK staining, included controls (Tj>EGFRi), and clarified limitations regarding MAPK signaling. Due to lethality, we could not perform an EGFR gain-of-function rescue. We also validated increased Hh signaling via qPCR and a Tj>UAS Ci control (Figure S4).

(5) Conceptual & Terminological Refinements: We revised our discussion of BMP signaling, ROS gradients, and testis-specific terminology. All figures and labels now accurately represent GSC scoring (single Vasa⁺ cells in contact with the niche).

(6) Figure & Methods Improvements: We enhanced image resolution, provided grayscale versions where needed,and expanded Materials & Methods to clarify experimental conditions.

These revisions strengthen our conclusions and address the reviewer’s concerns, ensuring a more precise and transparent presentation of our findings. To align with the reviewer’s s we have changed the title of the manuscript to “Superoxide Dismutases maintain niche homeostasis in stem cell populations”.

Specific responses to the reviewer’s comments:

(1) Data

a. Problems proving which mitochondria are associated with which lineage.

We acknowledge the challenge of distinguishing CySCs from GSCs without additional cell surface markers. To enhance clarity, we have incorporated the membrane marker Discs-large (Dlg) in our revised Figure 1 to better delineate cell boundaries, providing a clearer depiction of mitochondrial distribution in GSCs and CySCs.

b.There is no evidence that ROS diffuses from CySCs into GSCs.

We acknowledge the reviewer’s concern. There are reports which talks about diffusion of ROS across cells on which we have included a few lines in the discussion (line no. 274-276). We do understand that our previous quantifications showed ROS diffusion from CySC to GSC rather indirectly. Therefore, in revised manuscript we have measured ROS separately in the two cell populations. We found that the CySCs show higher ROS profile than GSCs (Fig 1R).

c.The changes in GST-GFP (redox readout) are possibly seen in differentiating germ cells (i.e., spermatogonia) but not in GSCs. This weakens their model that ROS in CySC is transferred to GSCs.

Thank you for your observation. We acknowledge that the changes in gstD-GFP (redox readout) are more prominent in differentiating germ cells. It is known that differentiating cells show higher ROS profile than the stem cells. Hence, expectedly the intensity of gstDGFP was lesser in stem cell zone compared to the differentiating zone. In our manuscript we are focussed on the redox state among stem cell populations. Therefore, we have included better quality images and measured the gstD1-GFP intensity individually in GSCs and CySCs (Figure 1R) by demarcating the cell boundaries (Figure 1M, S1C-D’’’). We found that CySCs show higher ROS profile than GSCs and enhancement of ROS in CySC by Sod1 depletion resulted in a consequent increase in ROS in GSCs. We believe this revision strengthens our model by addressing the potential discrepancy and providing a more comprehensive understanding of ROS dynamics within the GSC niche.

d.Most of the paper examines the effect of SOD depletion (which should increase ROS) on the CySC lineage and GSC lineage. One big caveat is that Tj-Gal4 is expressed in hub cells (Fairchild, 2016), so the loss of SOD from hub cells may also contribute to the phenotype. In fact, the niche in Figure 2D looks larger than the niche in the control in Figure 2C, arguing that the expression of Tj in niche cells may be contributing to the phenotype. The authors need to better characterize the niche in tj>SOD-RNAi testes.

We appreciate the reviewer’s insightful regarding the potential contribution of hub cell to the observed phenotype. We acknowledge that Tj-Gal4 is expressed in hub cells and this could influence the niche size and overall phenotype.

To address this concern, we have included an additional control using C587-Gal4, a CySC specific driver, to distinguish CySC-specific effects from potential hub contributions. All the effects on cell number observed in Tj>Sod1i was replicated in C587>Sod1i testis, except that the observed phenotypes were comparatively weaker. These indicate partial contribution of hub cells to the observed phenotype, exacerbating its severity. However, the effect of Sod1 depletion in CySC on GSC lineages remains significant. These findings have been incorporated into Figure S1- G,H,M and N) and incorporated in the discussion (line no.308311).

e. The Tj>SOD1-RNAi phenotype is an expansion of the Zfh1<sup+ CySC pool, expansion of the Tj+ Zfh1- cyst cells (both due to increased somatic proliferation) and a non-autonomous disruption of the germline.

We appreciate the reviewer’s observation. Our data confirm that Tj>SOD-RNAi leads to an expansion of both Zfh1<sup+ CySCs and Tj+ Zfh1- cyst cells, which we attribute to increased somatic proliferation. Additionally, we observe a non-autonomous disruption of the germline, likely due to dysregulated signaling from the altered somatic niche.

f. I am not convinced that MAPK signaling is decreased in tj>SOD-i testes. Not only is this antibody finicky, but the authors don't have any follow-up experiments to see if they can restore SOD-depleted CySCs by expressing an EGFR gain of function. Additionally, reduced EGFR activity causes fewer somatic cells (not more) (Amoyel, 2016) and also inhibits abscission between GSCs and gonial blasts (Lenhart 2015), which causes interconnected cysts of 8- to 16 germ cells with one GSC emanating from the hub.

We acknowledge that the dpERK antibody can be challenging. We took necessary precautions, including optimizing staining conditions and using positive control (Tj>EGFRi) (Figure: S4B). Our results consistently showed a decrease in dpERK levels in Tj>Sod1i testes, supporting our conclusion.

We agree that inclusion of an experiment using EGFR gain-of-function to rescue the effects of CySC-Sod1 depletion would have strengthened our findings. We had attempted this experiment; however, the progenies constitutively expressing EGFR under Sod1RNAi background were lethal, preventing us from completing the analysis.

We agree that our observations do not align with the reported effects of EGFR signaling on somatic cell numbers and abscission and we appreciate the references provided. Based on our observations, we feel that modulation of MAPK signaling in the niche probably, happens in a context-dependent manner. One possible explanation is that, the altered GSC -CySC balance and loss of contact in Tj>Sod1i testes, leads to insufficient ligand response, thereby failing to activate EGFR signaling. While it is well established that ROS can enhance EGFR signaling to promote cellular proliferation and early differentiation, our results indicate a more nuanced regulation in this context. However, further detailed analysis is required to completely understand the regulatory controls. We have clarified this point in the manuscript (line no.

313-320).

g. The increase in Hh signaling in SOD-depleted CySCs would increase their competitiveness against GSCs and GSCs would be lost (Amoyel 2014). The authors need to validate that Hh protein expression is indeed increased in SOD-depleted CySCs/cyst cells and which cells are producing this Hh. Normally, only hub cells produce Hh (Michel,2012; Amoyel 2013) to promote self-renewal in CySCs.

We appreciate the reviewer’s suggestion regarding the validation of Hh protein expression and its source. Since Tj-Gal4 is expressed in the hub, it is likely activating the Hh pathway and promoting CySC proliferation. Unfortunately, we could not procure Hh antibody to directly assess its protein levels. However, to address this, we performed real-time PCR from RNA derived from the tip region and found a significant increase in hh mRNA levels in SOD-depleted cyst cells. These findings support our hypothesis that elevated Hh signaling enhances CySC competitiveness, leading to GSC loss. To support this idea, we have included a Tj>Ci positive control which caused abnormal proliferation of Tj+ cells resulted in ablation of GSCs. We have incorporated these results in the revised manuscript (Results section, Figure S-4).

h.The increase in p4E-BP is an indication that Tor signaling is increased, but an increase in Tor in the CySC lineage does not significantly affect the number of CySCs or cyst cells (Chen, 2021). So again I am not sure how increased Tor factors into their phenotype.

We acknowledge the reviewer’s concern regarding the role of increased Tor signaling in our phenotype. The observed increase in Tor could indeed be a downstream effect of elevated ROS levels. However, establishing a direct causal relationship between Sod1 and Tor would require additional experiments, which we feel might be a good study in its own merit. To maintain clarity and focus in the revised manuscript, we have opted not to include this preliminary data at this stage.

I.The over-expression of SOD in CySCs part is incomplete. The authors would need to monitor ROS in these testes. They would also need to examine with tj>SOD affects the size of the hub.

We value the reviewer's . To address this, we have now monitored ROS levels in the testes upon SOD overexpression in CySCs using DHE (Figure S5 I). Our results indicate a significant reduction in ROS levels compared to controls.

Additionally, we examined hub size upon Sod1 overexpression and observed a slight, but statistically insignificant, reduction. As our study primarily focuses on ROS-mediated GSCCySC interactions, we did not include a detailed investigation on hub size regulation.

(2) Concept

Why would it be important to have a redox gradient across adjacent cells? The authors mention that ROS can be passed between cells, but it would be helpful for them to provide more details about where this has been documented to occur and what biological functions ROS transfer regulates.

We thank the reviewer for this insightful . We acknowledge that the concept of a redox gradient was not adequately conveyed, as the cell boundary was not clearly defined. To address this, we have revised our interpretation to propose that high ROS levels in one cell may influence the ROS levels in an adjacent cell through either direct transfer or as a secondary effect of altered niche maintenance signaling, rather than through the establishment of a gradient.

Regarding ROS transfer between cells, it has been documented in several biological contexts. For instance, hydrogen peroxide (H2O2) can diffuse through aquaporins, influencing signaling pathways in neighbouring cells (PMID: 17105724). We have incorporated these details and relevant references into the revised manuscript to enhance the conceptual understanding of ROS transfer.

(3) Issues with the scholarship of the testis

a. Line 82 - There is no mention of BMPs, which are the only GSC-self-renewal signal. Upd/Jak/STAT is required for the adhesion of GSCs to the niche but not self-renewal (Leatherman and Dinardo, 2008, 2010). The author should read a review about the testis. I suggest Greenspan et al 2015. The scholarship of the testis should be improved.

We appreciate the reviewer’s feedback regarding the role of BMPs in GSC selfrenewal, we have added this in the revised manuscript (line no. 83) We have now incorporated a discussion on BMP signaling as the primary self-renewal signal for GSCs, distinguishing it from the role of Upd/JAK/STAT in niche adhesion, as highlighted in Leatherman and Dinardo (2010). Additionally, we have cited and reviewed the work by Greenspan et al. (2015) and ensure a more comprehensive discussion of GSC regulation. These revisions can be found in the line no. 285-289 of the revised manuscript.

b. Line 82-84 - BMPs are produced by both hub cells and CySCs. BMP signaling in GSCs represses bam. So it is not technically correct to say the CySCs repress bam expression in GSCs.

We acknowledge the reviewer’s clarification regarding BMP signaling and its role in repressing bam expression in GSCs. We have revised the relevant section (line no.83-85).

c.Throughout the figures the authors score Vasa+ cells for GSCs. This is technically not correct. What they are counting is single, Vasa+ cells in contact with the niche. All graphs should be updated with the label "GSCs" on the Y-axis.

We appreciate the reviewer’s careful assessment of our methodology. We acknowledge that scoring Vasa⁺ cells alone does not definitively identify GSCs. Our quantification specifically considers single Vasa⁺ cells in direct contact with the niche. To ensure clarity and accuracy, we have updated all figure legends and Y-axis labels in the relevant graphs to explicitly state "GSCs" instead of "Vasa⁺ cells."

(4) Issues with the text

a. Line 1: multi-lineage is not correct. Multi-lineage refers to stem cells that produce multiple types of daughter cells. GSCs produce only one type of offspring and CySCs produce only one type of offspring. So both are uni-lineage. Please change accordingly.

We acknowledge the incorrect usage of "multi-lineage" and agree that both GSCs and CySCs are uni-lineage, as they each produce only one type of offspring. We have revised Line 1 accordingly and also updated the title.

b. Lines 62-75 - Intestinal stem cells have constitutively high ROS (Jaspar lab paper), so low ROS in stem cell cells is not an absolute.

We appreciate the clarification. We have revised Lines 62–75 to acknowledge that low ROS is not universal in stem cells, citing the Jaspar lab study on intestinal stem cells (Line 70). Thank you for the valuable insight.

c. Line 79: The term cystic is not used in the Drosophila testis. There are cyst stem cells (CySCs) that produce cyst cells. Please revise.

We have revised the text to replace "cystic" with the correct terminology, referring to cyst stem cells (CySCs) in the manuscript.

d. Line 90 - perfectly balanced is an overstatement and should be toned down.

Thank you for the suggestion. We have revised it to “balanced” instead of "perfectly balanced."

e. Line 98 - division of labour is not supported by the data and should be rephrased.

Thank you for the feedback. We have rephrased it (line no. 98-101) to avoid the term "division of labor".

f. Line 200 - the authors provide no data on BMPs - the GSC self-renewal cue - so they should avoid discussing an absence of self-renewal cues.

We appreciate the reviewer’s point. We have revised it to avoid discussing the absence of self-renewal cues, given that we do not present data on BMP signaling. This ensures that our conclusions remain within the scope of the provided data.

(5) Issues with the figures

a The images are too small to appreciate the location of mitochondria in GSCs and CySCs.

b. Figure 1

c. cell membranes are not marked, reducing the precision of assigning mitochondria to GSC or CySCs. It would be very helpful if the authors depleted ATP5A from GSCs and showed that the puncta are reduced in these cells, and did a similar set of experiments for the Tj-Gal4 lineage. It would also be very helpful if the authors expressed membrane markers (like myrGFP) in the GSC and then in the CySC lineage and then stained with ATP5A. This would pinpoint in which cells ATP5A immunoreactivity is occurring.

d. The presumed changes in gst-GFP (redox readout) are possibly seen in differentiating germ cells (i.e.,spermatogonia) but not in GSC. iii. Panels F, Q, and S are not explained and currently are irrelevant.

e. Figure 3K - The evidence to support less Ecad in GSCs in tj>SOD-i testes is not compelling as the figure is too small and the insets show changes in Ecad in somatic cells, not GSC. d. Figure 4:

f. Panel A, B The apparent decline (not quantified) may not contribute to the phenotype.

ii.dpERK is a finicky antibody and the authors are showing a single example of each genotype. This is an important experiment because the authors are going to use it to conclude that MAPK is decreased in the tj>SOD-i samples. However, the authors don't have any positive (dominantactive EGFR) or negative (tj>mapk-i). As is standing, the data is not compelling. The graph in F does not convey any useful information.

g. Figure S1D - cannot discern green on black. It is critical for the authors to show monochromes (grayscale) for thereabouts that they want to emphasize. I cannot see the green on black in Figure S1D.

h. Figure S4 - there is no quantification of the number of Tj cells in K-N.

We appreciate your detailed feedback regarding the figures in our manuscript. Below, we address each concern and outline the revisions we have made.

(a) Image Size and Mitochondrial Localization in GSCs and CySCs

We acknowledge the need for larger images to better visualize mitochondrial localization. We have now increased the resolution and size of the images in Figure 1. Additionally, we have included high-magnification insets to enhance clarity (Figure 1 B#)

(b) Figure 1 B,B#,C

(i) We have now marked cell membranes using Dlg to improve the precision of mitochondrial assignment to GSCs and CySCs and then stained for ATP5A, which clearly demarcates ATP5A immunoreactivity in specific cell types.

(ii) We have revisited the gstD-GFP (redox readout) data and now provide revised images (Figure S1C-D’’’) and quantification (Figure 1 R,S) to better illustrate changes in the redox state. It is indeed intense in differentiating germ cells as expected but also present in the stem cell zone.

(iii) Panels F, Q, and S have now been removed in the revised figure legend.

(C) Figure 3K: We have digitally magnified the figure size and improved contrast to better visualize E-cadherin levels. The insets have been revised to ensure they focus specifically on GSCs rather than somatic cells. Earlier, we quantified the E-cadherin intensity changes in the GSC-hub interface and provided statistical analysis to support our findings (Figure 3M).

(d) Figure 4: (i) Panels A and B have now been quantified, and we provide statistical comparisons to support our observations. (ii) We acknowledge the variability of dpERK staining. To strengthen our conclusions, we have provided negative (Tj>MAPK-i) controls (Figure S4 B). Additionally, we have removed panel F (MAPK area cover) to avoid confusion.

(e) We appreciate the suggestion regarding grayscale images and have provided the monochrome images for mitochondria and gstD-GFP image representation. We have now removed Figure S1D as it was no longer required.

(f) Figure S4: The quantification of the number of Tj-positive cells was actually included in the main figure along with statistical analysis.

(g) We sincerely appreciate the reviewer’s insightful s, which have significantly improved the quality and clarity of our manuscript. We hope that our revisions adequately address the concerns raised.

(6) Issues with Methods

a. Materials and Methods are not described in sufficient depth - please revise.

b. Note that Tj-Gal4 has real-time expression in hub cells and this is not considered by the authors. The ideal genotype for targeting CySCs is Tj-Gal4, Gal80TS, hh-Gal80. Additionally, the authors do not mention whether they are depleting throughout development into adulthood or only in adults. If the latter, then they must have used a temperature shift, growing the flies at 18C and then upshifting to 25C or 29C during adult stages.

c. The authors need to show data points in all of the graphs. Some graphs do this but others do not.

d. The authors state that all data points are from three biological replicates. This is not sufficient for GSC and CySC counts. Most labs count GSCs and CySCs from at least 10 testes of the correct genotype.

We appreciate the reviewer’s valuable feedback and have made the necessary revisions to improve the clarity and rigor of our study. Below, we address each concern in detail:

Materials and Methods

We have revised the Materials and Methods section to provide a more detailed description of the experimental procedures, including genotypes, sample preparation, and quantification methods.

Tj-Gal4 Expression and Experimental Design

We acknowledge the reviewer’s point regarding Tj-Gal4 expression in hub cells. While Tj-Gal4 is active in hub cells, our focus was on CySCs, and we have now included a discussion of this caveat in the revised manuscript (line no. 308-311)

Thank you for your suggestion on the ideal genotype for targeting CySCs. While we attempted to procure hh-Gal80, we couldn’t manage to get it, so we opted for another well-established Gal4 driver, C-587 Gal4, to target CySCs. Our results indicate that although the phenotypic changes are consistent across both drivers, the effects are significantly stronger with Tj-Gal4, highlighting the role of CySCs in this process with partial contributions from the hub. These findings have been incorporated into the revised manuscript (lines 309–311).

We now clarify whether gene depletion was conducted throughout development or restricted to adulthood. For adult-specific depletion using the UAS-Gal4 system, crosses were set up at 25°C, and after two days, progenies were shifted to 29°C and aged for 3–5 days at 29°C. This process is now explicitly detailed in the revised Methods section (line no. 345-348).

Data Presentation in Graphs

We have updated all graphs to ensure that individual data points are shown consistently across all figures.

Sample Size for GSC and CySC Counts

We acknowledge the reviewer’s concern regarding biological replicates. Our initial study was based on 10 biological replicates, each set consisting of at least 7-8 testes per genotype, in line with standard practice in the field. This change is reflected in the revised Results and Methods sections.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Reactive oxygen species (ROS), predominantly derived from mitochondrial respiratory complexes, have evolved as key molecules influencing cell fate decisions like maintenance and differentiation. These redox-dependent events are mainly considered to be cell intrinsic in nature, on contrary our observations indicate involvement of these oxygen-derived entities as intercellular communicating agents. In multi-lineage Drosophila germline, neighbouring Germline Stem Cells (GSCs) and Cystic Stem Cells (CySCs) maintain differential redox thresholds where CySCs by virtue of their higher redox-state regulate physiological ROS levels of germline. Disruption of the intercellular redox equilibrium between the two adjoining stem cell populations results in deregulated niche architecture and loss of GSCs, which was mainly attributed to loss of contact-based receptions and uncontrolled CySC proliferation due to ROS-mediated activation of self-renewing signals. Our observations hint towards the crucial role of intercellular redox gradients originating from somatic progenitors, CySCs in niche stability where they function not only as a source of their own maintenance cues but also serve as non-autonomous redox moderators of germline immortality. Our findings underscore the complexity of niche homeostasis and predicate the importance of intercellular redox communication in understanding stem cell microenvironments.

Revised version: Reactive oxygen species (ROS), predominantly derived from mitochondrial respiratory complexes, have emerged as key molecules influencing cell fate decisions like maintenance and differentiation. These redox-dependent events are mainly considered to be cell intrinsic in nature, on the contrary our observations indicate involvement of these oxygen-derived entities as intercellular communicating agents. In Drosophila male germline, Germline Stem Cells (GSCs) and neighbouring Cyst Stem Cells (CySCs) maintain differential redox thresholds where CySC have higher redox-state compared to the adjacent GSCs. Disruption of the redox equilibrium between the two adjoining stem cell populations by depleting Superoxide Dismutases (SODs) especially Sod1 results in deregulated niche architecture and loss of GSCs, which was mainly attributed to loss of contact-based receptions and uncontrolled CySC proliferation due to ROS-mediated activation of self-renewing signals. Our observations hint towards the crucial role of differential redox states where CySCs containing higher ROS function not only as a source of their own maintenance cues but also serve as non-autonomous redox moderators of GSCs. Our findings underscore the complexity of niche homeostasis and predicate the importance of intercellular redox communication in understanding stem cell microenvironments.

Introduction

Earlier version: Studies from the past few decades have shown the apparent role of ROS in influencing various biological processes1–3. Generally, the sources of ROS are compartmentalized in co-occurrence with their targets mediating important signalling events4. Mitochondria, a major source of oxidant species5,6, localize dynamically towards nucleus, effecting oxidation induced reshaping of gene expression profiles7. Localized ROS production are contributed by NADH oxidases distributed across different cellular regions8. Intracellular hydrogen peroxide gradients maintained by thioredoxin system allow intercompartmental exchanges among endoplasmic reticulum (ER), mitochondria and peroxisomes9–11.

Revised version: Studies from the past few decades have shown the apparent role of ROS in influencing various biological processes1-3. ROS are usually produced in specific cellular compartments close to their target molecules to regulate important signaling pathways.4. Mitochondria, a major source of oxidant species5,6, localize dynamically towards nucleus, effecting oxidation induced reshaping of gene expression profiles7. Localized ROS production are contributed by NADH oxidases distributed across different cellular regions8. Intracellular hydrogen peroxide gradients maintained by thioredoxin system allow intercompartmental exchanges among endoplasmic reticulum (ER), mitochondria and peroxisomes9-11.

Earlier version: The process of redox relays is conserved and plays a fundamental role in self-renewal and differentiation of stem cell populations12. Stem cells whether embryonic or adult maintain a low redox profile, characterized by subdued mitochondrial respiration13–17. Levels of ROS are tightly regulated and elevated amounts promote early differentiation and atypical stem cell behaviour18. However, leading evidences suggest that many transcription factors require oxidative environment for maintenance of pluripotent states. For instance, physiological ROS levels play a crucial role in genome maintenance of embryonic stem cells13. Multipotent hematopoietic progenitors and neural stem cells require relatively high baseline redox for their maintenance14,16,19. Suppression of Nox system or mitochondrial ROS compromises the maintenance of these self-renewing populations and promote their differentiation or death18. These evidences point towards essentiality of a perfectly tuned redox state for balancing the pluripotent and differentiated states. But how the stem cells harmonize their redox potential by possessing a restrained oxidant system is not very clear.

Revised version: The process of redox relays is conserved and plays a fundamental role in self-renewal and differentiation of stem cell populations12. Stem cells whether embryonic or adult maintain a low redox profile, characterized by subdued mitochondrial respiration13-17. Levels of ROS are tightly regulated and elevated amounts promote early differentiation and atypical stem cell behaviour18. However, leading evidences suggest that many transcription factors require oxidative environment for the maintenance of pluripotent states. For instance, physiological ROS levels play a crucial role in genome maintenance of embryonic stem cells13. Multipotent hematopoietic progenitors, intestinal stem cells 19 and neural stem cells require relatively high baseline redox for their maintenance14,16,20. Suppression of Nox system or mitochondrial ROS compromises the maintenance of these self-renewing populations and promote their differentiation or death18. These evidences point towards the essentiality of a well-tuned redox state for balancing the pluripotent and differentiated states. However, how stem cells regulate their redox potential by possessing a restrained oxidant system is not very clear.

Earlier version: We addressed this fundamental question in multi-stem cell lineage-based niche architecture in Drosophila testis. The testicular stem cell niche is composed of a central cluster of somatic cells called the hub which contacts eight to eleven GSCs arranged in a round array20,21. A pair of cystic stem cells (CySCs) originating the somatic lineage cyst cells, enclose each GSC and make their independent connections with the hub and GSCs via adherens junctions22–26. The hub cells secrete Unpaired family of cytokines that defines niche occupancy and self-renewal of both GSCs and CySCs through Jak-Stat signalling27,28. CySCs also repress GSC differentiation by suppressing transcription of bag-of-marbles (Bam)29,30. Concerted asymmetric division of both GSC and CySC is essential for proper cyst formation31. A developing cyst contains a dividing and differentiating gonialblast encircled by cyst cells that provide nourishment to developing spermatids32,33. However, the maintenance of GSCs in the spatially controlled microenvironment is still not very clear.

Revised version: We addressed this fundamental question in two-stem cell population-based niche architecture in Drosophila testis. The testicular stem cell niche is composed of a central cluster of somatic cells called the hub which contacts eight to eleven GSCs arranged in a round array21,22. A pair of cyst stem cells (CySCs) enclose each GSC and make their independent connections with the hub and GSCs via adherens junctions23-28. The hub cells secrete self-renewal factors essential for both GSCs and CySCs maintenance involving signalling cascades like Jak-Stat signalling29,30, BMP (Bone Morphogenetic Signalling)31, Hedgehog signalling32. The hub and CySCs produce BMPs which repress GSC differentiation by suppressing transcription of bag-of-marbles (Bam)31,33. Concerted asymmetric division of both GSC and CySC is essential for proper cyst formation34. A developing cyst contains a dividing and differentiating gonialblast encircled by cyst cells that provide nourishment to developing spermatids35,36. However, the maintenance of GSCs in the spatially controlled microenvironment is still not very clear.

Earlier version: We found the existence of a perfectly balanced redox differential between GSC and CySC essential for niche homeostasis. CySCs by virtue of their higher redox threshold and clustering of mitochondria at GSC-CySC interface, generate an intercellular redox gradient that maintain physiological ROS levels in GSCs. Alterations in CySC redox state affected self-renewal and differentiation propensities of the GSCs. Intercellular redox-imbalance induced niche disequilibrium through precocious differentiation of GSCs and deregulated CySC proliferation, due to activation of pro-proliferative transductions and loss of cell-cell contact signalling. Our results in multi-stem cell lineage-based germline system indicate a nuanced division of labour where a higher redox state in CySCs not only maintains their self-renewing processes but also non-autonomously foster GSCs maintenance.

Revised version: We found the existence of a balanced differential redox state between GSC and CySC essential for niche homeostasis. CySCs by virtue of their higher redox threshold and more clustered mitochondria generate an intercellular redox state that affects the physiological ROS levels in GSCs. Alterations in CySC redox state resulted in self-renewal and differentiation propensities of the GSCs. Intercellular redox-imbalance induced niche disequilibrium through precocious differentiation of GSCs and deregulated CySC proliferation, due to activation of pro-proliferative transductions and loss of cell-cell contact signalling. Our results indicate a sophisticated interplay in these two stem cell populations where a higher redox state in CySCs not only supports their self-renewing processes but also non-autonomously promotes the maintenance of GSCs.

Results

Added in the revised version: CySCs maintain higher differential ROS in comparison to adjacent GSCs

Removed in the revised version: Redox state of GSC is maintained by asymmetric distribution of CySC mitochondria

Earlier version: The apical region of the Drosophila testis incorporates a cluster of differentiated cells the hub surrounded by GSCs in a rosette arrangement with each GSC being enclosed by two CySCs (Fig. 1A). The mitochondria morphology and distribution in this cluster was checked by immunostaining ATP5A subunit of mitochondrial ATPase. We observed fragmented mitochondria in Vasa+ GSCs as compared to a fused organization in Traffic jam (Tj)+ CySCs (Fig. 1B#, C, D). Distribution profiles of mitochondria showed its higher propensity to overlap with Tj, indicating CySCs to be a mitochondria dense region compared to GSC (Fig. 1F); and its mitochondria to be predominantly clustered near the GSC-CySC boundary, illustrated by intensity profile (Fig. S1A’’’) and X-Z and Y-Z sectioning (Fig. S1B-B’). The ATP5A labelling overlapped with TFAM-GFP, a mitochondrial transcription factor34, further confirming the patterning observed between these two stem cell populations (Fig. S1C’).

Revised version: The apical region of the Drosophila testis incorporates a cluster of differentiated cells, the hub surrounded by GSCs in a rosette arrangement with each GSC being enclosed by two CySCs (Fig. 1A). Immunostaining ATP5A subunit of mitochondrial ATPase and marking cell boundary by discs-large (Dlg) indicated different mitochondria morphology (Fig. 1G) and distribution among these two stem cell populations. In wild type adult testis, we observed fragmented mitochondria in Vasa+ GSCs as compared to a fused organization in CySCs (Fig. 1B#, C, D, and S1B-B’’). The number of mitochondria per cell was higher in CySC (Fig 1F) and had a more elongated shape (reduced circularity) than the GSCs (Fig 1G). The ATP5A labelling overlapped with TFAM-GFP, a mitochondrial transcription factor37, further confirming the patterning observed between these two stem cell populations (Fig. S1A). Mitochondria localization and dispersion pattern in GSCs and CySCs, corresponded with the intensity variance of the ROS reporter line gstD1-GFP38 at the niche (Fig. 1L and S1C’’’). CySCs exhibited a higher baseline level of ROS compared to GSCs (Fig.1R).

Mitochondrial Distribution and Redox Profile of GSCs and CySCs in Adult Drosophila Testes

Earlier version: (A) Schematic representation of adult Drosophila testicular niche showing arrangement of different stem cell populations (GSC – Germinal Stem Cell, CySC – Cystic Stem Cell). (B-E) Differential distribution of mitochondria labelled with ATP5A (green) in Vasa+ GSCs and Tj+ CySCs of wild-type fly testis; B# shows digitally zoomed image of the dotted area. (F) Relative mitochondria density profile of GSCs and CySCs denoted through cooccurrence of ATP5A with Tj or Vasa. A.U. – arbitrary units. (G-P) Redox profiling of testicular stem cell niche using gstD1-GFP as intrinsic ROS reporter in control (G-K) and Tj-Gal4 driven Sod1RNAi (Sod1i) testis (L-P). (Q) Gaussian distribution curve for the variations in gstD1-GFP fluorescence at CySC-GSC intersection region (inset) in Tj-Gal4 control and Sod1i testis. (R) Quantification of mean gstD1-GFP fluorescence intensity (MFI) upon Sod1 knockdown represented as fold change over control. Data denotes mean ± s.e.m, n = 25 (S). Colocalization intensity between Vasa and gstD1-GFP expressed through Manders’ co-localization coefficient (MCC) among control and Sod1i cells (G). Controls are the indicated driver line crossed with Oregon R+. Scale bar – 10 µm. See also Figure S1.

Revised version: (A) Schematic representation of adult Drosophila testicular niche showing arrangement of different stem cell populations (GSC – Germinal Stem Cell, CySC – Cyst Stem Cell). (B-E) Differential distribution of mitochondria labelled with ATP5A (monochrome/green) in Vasa+ GSCs of wild-type fly testis with cellular boundaries marked by Dlg; B# shows digitally zoomed image of the dotted area. (F) Quantification of mitochondrial size per cell. Bars represent mean ± s.e.m., n = 10 fields of view; ***P (unpaired t-test)<0.0001. (G) Representation of mitochondrial shape,. that is circular or elongated (less circular) in two stem cell populations. Bars represent mean ± s.e.m., n = 10 fields of view; ***P (unpaired t-test)<0.0001. (H-Q) Redox profiling of testicular stem cell niche using gstD1-GFP as intrinsic ROS reporter in control (H-L) and Tj-Gal4 driven Sod1RNAi (Sod1i) testis (M-Q). (R) Quantification of gstD1-GFP mean fluorescence intensity (MFI) represented as fold change between GSCs and CySCs in control. Data denotes mean ± s.e.m,, n = 25, ***P(unpaired t-test)<0.0001. (S) Quantification of cell-specific ROS content. Data denotes mean ± s.e.m, n = 25, ***P(unpaired t-test)<0.0001. Controls are the indicated driver line crossed with Oregon R+. Scale bar - 10 µm. See also Figure S1.

Removed in the revised version: Mitochondria localization near the GSC-CySC boundary, together with its dispersion pattern was in accordance with the intensity variance of ROS reporter line gstD1-GFP 35 at the niche (Fig. 1H and S1D-D’). The redox gradient intensifies at the Tj-Vasa boundary and then gradually diffuses into Vasa+ GSCs, suggesting an establishment of GSC redox state through CySCs (Fig. 1Q; left panel). A GSC intrinsic oxidative environment should have resulted in a constant continuum of gstD1-GFP intensity across the Vasa-stained structure. We confirmed the greater ROS gradient at the GSC-CySC interface by 2D-Plot profile of Vasa, Tj and gstD1-GFP across the different cell types (Fig. S1E).

Earlier version: To affirm the role of CySC in influencing GSC redox state, we asked if disrupting dismutases in CySCs would influence the redox state of GSCs. We depleted the expression of superoxide dismutase 1 (Sod1) in CySCs and early differentiating cyst cells using Tj-Gal4 driver (Tj>Sod1i). Higher CySC redox state resulted in corresponding increase in GSC ROS profile, as indicated by a greater degree gstD1-GFP intensity and DHE fluorescence at the Vasa-gstD1GFP (VՈG) intersection (Fig. 1K, P, R and S1F), suggesting non-autonomous maintenance of GSC redox gradient by CySCs (Fig. 1Q; right panel). The surge in GSC ROS levels upon lowering of antioxidant defences in CySCs, was additionally confirmed through enhanced Vasa:gstD1-GFP colocalization intensities in Sod1i populations (MCCmean = 0.482) with respect of controls (MCCmean = 0.276) (Fig. 1S). These observations hinted towards a potential role of CySCs and its GSC polarized mitochondria in elevating the cellular ROS gradient at the interface between two stem cell populations.

Revised version: To affirm the role of CySC in influencing GSC redox state, we asked if disrupting the dismutases in CySCs would influence the redox state of GSCs. We depleted the expression of superoxide dismutase 1 (Sod1) in the hub, CySCs and early differentiating cyst cells (CCs) using Tj-Gal439 driver (Tj>Sod1i) that resulted in a net increase in gstD1-GFP intensity (Fig. 1Q, D’’ and S1E). The elevated redox state in CySCs led to a corresponding increase in ROS levels in adjacent GSCs, as evidenced by the increased gstD1-GFP intensity (Fig.1Q and S) and an overall rise in superoxide levels, confirmed through DHE fluorescence analysis (Fig.S1F). This suggests a potential redox crosstalk between these stem cell populations, where the redox state of CySCs influences the oxidative status of neighbouring GSCs.

Balanced CySC redox profile is crucial for its proliferation and GSC maintenance

Added in the revised version: Alongside the differential redox profile, we observed that elevated ROS levels in Tj>Sod1i had a striking effect on the increase of DAPI+ nuclei at the testis tip (Fig. 2A’’). This overcrowding was attributed to both an increase in the number of Tj+ CySCs and early differentiating CCs, as well as their positional shift (Fig. 2A’, B’, and E). To further validate these findings, we used an alternative driver line, C-587-Gal4, which specifically drives expression in CySCs and CCs. This also resulted in an increase in the total number of Tj+ cells (Fig. S1H’, and M) although to a lesser extent than Tj-Gal4, suggesting that signals from the hub may also partially influence CySC proliferation. Given the stronger phenotypic effect observed with Tj-Gal4, we proceeded with this driver line for subsequent experiments.

Removed in the revised version: Along with the shift in the redox gradient, we observed that high ROS in Tj>Sod1i had a striking effect on the increase in DAPI+-nuclei at the tip of testis (Fig. 1G and L). This overcrowding was attributed to an increase in number and positional shift in Tj+ CySCs and early differentiating cyst cells (Fig. 2A’, B’ and E). The enhancement in the number of CySCs was also reflected by ∼2-fold change in Zfh1+ cells (Fig. 2F). Vasa+ cells flanking the hub however, substantially reduced in number (Fig. 2C’’-D’’ and G), which was validated through lower detectable levels of Vasa even when its expression per cell remained unchanged (Fig. S2O). The observations were further verified using Sod1i localized in different chromosome to avoid any chromosome or balancer-based biasness (Fig. S1G-L). However, the phenotypic effect of higher ROS was limited to its origin in CySCs only because ablation of GSC redox status did not cause any marked change in niche composition. Nos-Gal4 driven knock-down of Sod1 in GSCs did not result in significant change in CySC number (Fig. S2A’, B’ and E) but effected a considerable reduction on Vasa+ cells (Fig. S2C’, D’ and F), aligning with previous observations17. Although, Sod1 is the predominant enzyme which also localizes in the intermembrane space of mitochondria, we observed a similar result upon depleting the matrix localized Sod2 in both GSCs and CySCs (Fig. S2G-M), indicating involvement of mitochondrial ROS for the observed cellular phenotypes.

Redox disequilibrium in Cyst stem cell lineage deregulates early CySCs and decreases the GSC number

Earlier version: (A-B) Distribution of Tj+ CySC/early cyst cell around the hub (FasIII) in control and driven Sod1RNAi testis. (C-D) Represents the rosette arrangement of Vasa+ GSCs flanking the hub. The digitally magnified region around the hub (dotted line) as seen in control (C#) and Sod1RNAi (D#). (E-G) Mean number of Tj+ cells, Zfh1+ early CySCs and GSCs in control and Sod1RNAi lines shown as mean ± s.e.m, n = 10. (H) Construct design of the FUCCI reporter line for tracking the cell cycle stages in Drosophila tissues, containing degrons of Cyclin B and E2F1 proteins fused with RFP or GFP. G1, S and G2/M is represented by GFP+, RFP+ and dual labelled cells respectively. (I-J) Comparative changes in the total cell population (I) or Zfh1+ cells (J) present in G1, S and G2/M phase at the niche zone between control and Tj>Sod1i. Data points denotes mean ± s.e.m, n = 10. Scale bar – 10 µm. See also Figure S2.

Revised version: (A-B) Distribution of Tj+ CySC/early cyst cell around the hub (FasIII) in control and Tj driven Sod1RNAi testis. (C-D) Represents the rosette arrangement of Vasa+ GSCs flanking the hub. The digitally magnified region around the hub (dotted line) as seen in control (C#) and Sod1RNAi (D#). (E-G) Mean number of Tj+ cells, Zfh1+ early CySCs and GSCs in control and Sod1RNAi lines shown as mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001, **P(unpaired t-test)<0.0001. (H) Construct design of the FUCCI reporter line for tracking the cell cycle stages in Drosophila tissues, containing degrons of Cyclin B and E2F1 proteins fused with RFP or GFP. G1, S and G2/M is represented by GFP+, RFP+ and dual labelled cells respectively. (I-J) Comparative changes in the cell cycle phases (I) or Zfh1+ cells (J) present in G1, S and G2/M phase at the niche zone between control and Tj>Sod1i. Data denote mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001, **P(unpaired t-test)<0.001. (K) Mean number of pH3+ cells, a mitotic marker. Data points denote mean ± s.e.m, n = 10. Scale bar - 10 µm. See also Figure S2.

Added in the revised version: The enhancement in the number of CySCs was also reflected by ∼2-fold change in Zfh1+ cells (Fig. 2F, S2R-S). However, Vasa+ cells flanking the hub (GSCs) showed a substantial reduction in number (Fig. 2C’’-D’’ and G). This decline was further validated by western blot analysis, which revealed lower overall detectable levels of Vasa protein, indicating a potential impact of the altered redox state on GSC maintenance and stability. (Fig. S2N). C-587-Gal4 also showed a reduction in the number of neighbouring GSCs (Fig.S1N). The observations were further verified using Sod1i localized in different chromosome to avoid any chromosome or balancer-based biasness (Fig. S1I-L and O-P). However, the phenotypic effect of higher ROS was limited to its origin in CySCs only because ablation of GSC redox status did not cause any marked change in niche composition. Nos-Gal4 driven knock-down of Sod1 in GSCs did not result in significant change in CySC number (Fig. S2A’, B’ and E) but effected a considerable reduction on Vasa+ cells (Fig. S2C’, D’ and F), aligning with previous observations17. Although Sod1 is the predominant enzyme which also localizes in the intermembrane space of mitochondria, we observed a similar result upon depleting the matrix localized Sod2 in both GSCs and CySCs (Fig. S2G-M), indicating the involvement of mitochondrial ROS for the observed cellular phenotypes.

Earlier version: To confirm active proliferation of cells and rule out the possibility of arrested growth, we utilized fly-Fluorescent Ubiquitination-based Cell Cycle Indicator (FUCCI) system36 which comprises of two reporter constructs marking G1/S transition (green), S (red) and G2/early mitosis (yellow) (Fig. 2H). Tj>Sod1i testes presented a ∼2-fold increase in cell number across all stages of cell-cycle (Fig 2I). The enhanced number of dividing cells was contributed mainly by progressive division of CySCs, showing >2-fold difference in accumulation of Zfh1+ nuclei in S and G2/M phases (Fig. 2J, S2PV-VII and QV-VII). The number of Zfh1 labelled cells in G1/S transition was also substantially more (Fig. S2PIII and QIII), indicating dysregulated redox ensued increased mitotic index of CySCs. Downregulation of Sod1 in GSCs through Nos-Gal4 demonstrated no significant change in G1/S, S and G2/M phase numbers, confirming our previous observation that altering ROS levels in GSCs does not have any non-autonomous effect on CySCs (Fig. S2N).

Revised version: To confirm the active proliferation of cells and rule out the possibility of arrested growth, we utilized fly-Fluorescent Ubiquitination-based Cell Cycle Indicator (FUCCI) system40 which comprises of two reporter constructs marking G1/S transition (green), S (red) and G2/early mitosis (yellow) (Fig. 2H). Testes from Tj>Sod1i exhibited an approximately two-fold increase in cell number across all stages of the cell cycle (Fig. S2O). However, the percentage of cells in each phase remained unchanged (Fig 2I), indicating that the observed increase in different phases is a consequence of more cells entering proliferation rather than alterations in the duration of different cell-cycle stages. The enhanced number of dividing cells was contributed mainly by progressive division of CySCs, showing >2-fold difference in the accumulation of Zfh1+ nuclei in S and G2/M phases (Fig. 2J, S2RV-VII and SV-VII). The number of Zfh1 labelled cells in G1/S transition was also substantially more (Fig. S2RIII and SIII), indicating dysregulated redox ensued an increased mitotic index of CySCs. The enhanced mitotic activity of these proliferating niche cells is further supported by increased phospho-histone H3 (PH3) incorporation, indicating a higher mitotic frequency compared to the control (Fig. S2T, U), along with elevated cyclin D expression (Fig. S2Q). Downregulation of Sod1 in GSCs through Nos-Gal4 demonstrated no significant change in G1/S, S and G2/M phase numbers, confirming our previous observation that altering ROS levels in GSCs does not have any non-autonomous effect on CySCs (Fig. S2P).

CySC induced ROS couples precocious differentiation of stem cell lineages

Added in the revised version: Since in Tj>Sod1i testes, the increased number of Tj+ cells, marking CySCs and early differentiating cells, exceeded the number of CySC-specific Zfh1+ cells (Fig. 2J and S2O), we checked for parallel enhancement of cellular differentiation using the corresponding marker, Eya (Eyes absent).

Earlier version: Since, in Sod1i testis the number of Tj+ CySCs and early differentiating cells37 were elevated and more than CySC-specific Zfh1+ cells (Fig. 2I and J), we sought to label these cells with another differentiation marker Eya (Eyes absent). We found that Eya+ cells clustered towards the hub along with a significant increase in their number in Sod1 depleted CySCs (Fig. 3A, B and D). The supposed premature expression of Eya resulted in a population of differentiating cyst precursor cells co-expressing Eya-Zfh1 (Fig. 3A’’, B’’ and D). This population was found to be deviating from the normal partitioning of wild type representative populations (Fig. 3C). On contrary, the same cells would not show a similar phenotypical arrangement when Sod1i depletion was driven in the germline lineage by Nos-Gal4 (Fig. S3A-D). These observations hint towards an obvious change in redox state of GSCs affecting their number, even when generated non-autonomously from CySCs.

Revised version: We found that Eya+ cells clustered towards the hub along with a significant increase in their number in Sod1 depleted CySCs (Fig. 3A, B and D). This increased number of Eya+ cells does not suggest that these differentiating cells are proliferative. Instead, we propose that the knockdown of Sod1 may alter the timing or regulation of cyst cell differentiation, leading to an accumulation of Eya+ cells near the niche. The supposed premature expression of Eya resulted in a population of differentiating cyst precursor cells co-expressing Eya-Zfh1 (Fig. 3A’’, B’’ and D). This population was found to be deviating from the normal partitioning of wild-type representative populations (Fig. 3C). In contrast, depletion of Sod1i in the germline lineage using Nos-Gal4 did not result in a similar phenotypic arrangement (Fig. S3A-D). These findings suggest that alterations in the redox state of GSCs can impact their numbers, even when the changes are induced non-autonomously through CySCs.

ROS imbalance in CySCs promotes differentiation of both GSCs and CySCs

Earlier version: (A-D) Number of early CySCs (Zfh1+) (A’-B’), late differentiating cyst cells (Eya+) (A-B) and Zfh1, Eya co-expressing population were imaged (A’’-B’’), represented schematically (C) and quantified with respect to control and Sod1RNAi lines (D). Data points represent mean ± s.e.m, n = 10. (E-F) The effect of Sod1 depletion in CySCs on the differentiation status of GSCs as observed using Bam-GFP reporter line. (G) The relative distance of the differentiation initiation zone (Bam+) from the hub (red) in control and Sod1i testis shown as mean ± s.e.m, n = 10. (H-J) The shape and size of the spectrosomes marked with α-spectrin (green) were imaged (I’’-J’’) and quantified for their distance from the hub (marked with asterisk) (H), n = 10. Branched fusome marks differentiating populations (I’-J’). (K-L) Comparative staining of CySC-GSC contacts through adherens junction using E-cadherin (monochrome/green). Dotted area near the hub has been expanded as inset to show loss of E-cadherin network. (M) 3D surface intensity histogram for E-cadherin expression near the hub. (N-O) Fold change in expression of Stat dependent transcripts Socs36E (N) and Ptp61F (O) among control and Sod1i niche. Data is shown as mean ± s.e.m, n = 3. Scale bar – 10 µm. See also Figure S3.

Revised version: (A-D) Number of early CySCs (Zfh1+) (A’-B’), late differentiating CCs (Eya+) (A - B) and Zfh1, Eya co-expressing population (A’’-B’’) were imaged, represented schematically (C) and quantified with respect to control and Sod1RNAi lines (D). Data points represent mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001. (E-F) The effect of Sod1 depletion in CySCs on the differentiation status of GSCs as observed using Bam-GFP reporter line. (G) The relative distance of the differentiation initiation zone (Bam+) from the hub (red) in control and Sod1i testis shown as mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001. (H-J) The shape and size of the spectrosomes marked with α-spectrin (green) were imaged (I’’-J’’) and quantified for their distance from the hub (marked with asterisk) (H), n = 10, ***P(unpaired t-test)<0.0001. Branched fusome marks differentiating populations (I’-J’). (K-L) Comparative staining of CySC-GSC contacts through adherens junction using E-cadherin (monochrome/green). Dotted area near the hub has been expanded as inset to show loss of E-cadherin network. (M) E-cadherin intensity histogram plot generated from ImageJ representing the mean intensity of expression in the region flanking the hub and GSCs. (N-O) Fold change in expression of Stat dependent transcripts Socs36E (N) and Ptp61F (O) among control and Sod1i niche, obtained through qPCR. Data is shown as mean ± s.e.m, n = 3, ***P(unpaired t-test)<0.0001. ns – not significant. Scale bar - 10 µm. See also Figure S3.

Earlier version: Together with reduction in Vasa+ GSCs shown earlier (Fig. 2G), we also detected gonialblast expressing differentiation-promoting factor Bam to be much closer to the hub in Tj>Sod1i testes (Fig. 3E’’ and F’’). The mean distance between the hub and Bam+ cells was found to be shortened by ∼7 microns (Fig 3G). The relative advancement of spermatogonial differentiation towards the hub was also ascertained by tracking the transformation of GSC-specific spectrosome into branched fusomes connecting the multi-cell gonialblast (Fig. 1A). The branching fusomes were found more proximal to hub than controls (Fig. 3H, I’’ and J’’), along with parallel reduction in number of spectrosomes, implying a decline in early-stage germ cells (Fig. S3F). This loss could possibly be attributed to a higher oxidizing environment in GSCs and loss of intercellular contacts, particularly with hub cells (Fig. 2C# and D#). The observation corroborated with decreased expression pattern of E-cadherin in Tj>Sod1i particularly near the hub (Fig. 3K-M and S3G) and could be one of the reasons for GSCs dissociation and differentiation (Fig. 2D#), due to absence of self-renewing signals. Indeed, disengagement of GSCs from the hub caused reduction in Socs36E and Ptp61F transcripts that act as drivers of E-cadherin expression and are induced by maintenance factor STAT originating from the hub (Fig. 3N and O). These results suggest that disruption of intercellular redox gradient effected premature differentiation of stem cells in the niche.

Revised version: Together with reduction in Vasa+ GSCs shown earlier (Fig. 2G), we also detected gonialblast expressing differentiation-promoting factor Bam to be much closer to the hub in Tj>Sod1i testes (Fig. 3E’’ and F’’). The mean distance between the hub and Bam+ cells was found to be shortened by ∼7 microns (Fig 3G). The relative advancement of spermatogonial differentiation towards the hub was also ascertained by tracking the transformation of GSC-specific spectrosome into branched fusomes connecting the multi-cell gonialblast (Fig. 1A). The branching fusomes were found more proximal to hub than controls (Fig. 3H, I’’ and J’’), along with parallel reduction in number of spectrosomes, implying a decline in early-stage germ cells (Fig. S3F). This loss could possibly be attributed to loss of intercellular contacts, particularly with hub cells (Fig. 2C# and D#). The observation corroborated with decreased expression pattern of E-cadherin in Tj>Sod1i flanking the hub and GSCs (Fig. 3K-M) and could be one of the reasons for GSCs dissociation and differentiation (Fig. 2D#), Disengagement of GSCs from the hub caused reduction in Socs36E and Ptp61F transcripts that act as drivers of E-cadherin expression (Fig. 3N and O). These results suggest that disruption of intercellular redox gradient affected the premature differentiation of stem cells in the niche.

Disrupted niche architecture compromises GSC-CySC communication

Earlier version: GSCs and CySCs intercommunicate through redox dependent factors such as EGFR, PI3K/Tor to support CySC maintenance38. EGF ligands secreted by GSCs maintain differences in cell polarity and segregates self-renewing from differentiating populations39. Tj>Sod1i cells showed lower levels of cell-polarity marker Dlg (Fig. 4A”, B” and S4A), that expanded more into the differentiated zone (Fig. S4B). Since, Dlg expression is dependent on EGFR40, impairment in its reception was reflected in net reduction in pErk levels across sections in Sod1i testis (Fig. S4C). Due to loss of contact, this reduction was mainly limited to regions surrounding the hub, in CySC zone (Fig. 4C’’-D’’ and E). Differentiated cyst cells presented relatively higher levels of pErk expression (Fig. 4E, S4D’), particularly in Tj>Sod1i tissues where the pErk area was significantly extended, mapping with that of Dlg (Fig. 4F) and located comparatively closure to the hub (Fig. S4D), supporting its redox-dependent enhancement to promote differentiation. The deregulated proliferation of cystic cells reconciled with higher levels of its self-renewing inducers PI3K effector, p4E-BP (Fig. 4G’, H’ and I) and Hedgehog (Hh) (Fig. 4M’ and N’) 41 in Tj>Sod1i testis. The noticeably high p4E-BP was associated with Cyclin-D accumulation (Fig. 4J). Similarly, depletion of Sod1 in CySC resulted in elevated expression of Hh receptor Patched (Ptc) (Fig. 4O) which extended to regions farther from the hub, overlapping with expanded Tj+ cells (Fig. 4M-N). Since, Ptc itself is a transcriptional Hh target, its expression in CySCs suggest active Hh signalling42, further supported by higher levels and similar pattern of Hh transcriptional effector Cubitus interruptus (Ci) in Sod1i testis (Fig. 4P and S4E’’’-F’’’). Since, Ptc, Ci and a GPCR-like signal transducer Smoothened (Smo) are all transcriptionally driven by Hh, and Hh itself is regulated by higher levels of Ci43,44, we expectedly found higher levels of transcripts for these pathway genes (Fig. S4G-J), indicating ROS-mediated upregulation of Hh activity. The overt proliferation of CySCs under high ROS conditions was suppressed by reducing gene dosage either by single copy of PI3KDN (dominant-negative allele) or Hh-RNAi (Fig. 4L and Fig. S4K-N) which was also associated with an equivalent maintenance of GSC population (Fig. 4K and S4O-R). Together these observations suggest that intercellular redox gradients balance the stem-cell populations in the niche through CySC maintenance factors.

Revised version: GSCs and CySCs intercommunicate through several factors such as EGFR, PI3K/Tor, Notch41 to support CySC maintenance42. EGF ligands secreted by spermatogonia maintain differences in cell polarity and segregates self-renewing from differentiating populations43. Tj>Sod1i cells showed lower levels of cell-polarity marker Dlg (Fig. 4A”, B” and E), that expanded more into the differentiated zone (Fig. S4A). Since, Dlg expression is dependent on EGFR44, impairment in its reception was reflected in net reduction in pErk levels across sections in Sod1i testis (Fig. 4C’, D’ and F). This reduction in overall Erk levels may be attributed to the loss of cell contact and altered GSC-CySC balance in Sod1i testis, probably resulting in lower EGF response. The uncontrolled proliferation of cyst cells correlates with increased levels of self-renewal inducers, including the elevated expression of Hedgehog (Hh) pathway components 32 in Tj>Sod1i testis. The depletion of Sod1 in CySC resulted in elevated expression of Hh receptor Patched (Ptc) (Fig. 4I, S4V) which extended to regions farther from the hub, overlapping with expanded Tj+ cells (Fig. 4G and H). Since, Ptc itself is a transcriptional Hh target, its expression in CySCs suggests active Hh signalling45, further supported by higher levels and a similar pattern of Hh transcriptional effector Cubitus interruptus (Ci) in Sod1i testis (Fig. 4M’’, N’’, J and S4W). Since, Ptc, Ci, and a GPCR-like signal transducer Smoothened (Smo) are all transcriptionally driven by Hh, and Hh itself is regulated by higher levels of Ci46,47, we expectedly found higher levels of transcripts for these pathway genes (Fig. S4C-F). The overt proliferation of CySCs under high ROS conditions was suppressed by reducing gene dosage using Hh-RNAi (Fig. 4L and Fig. S4G-I), which was associated with rescue in the GSC population (Fig. 4K and S4J-L). The levels of Sod1 transcripts under single and Sod1, hh double depletion were almost similar (Fig. S4M). To determine whether the changes observed in the germ cell population were a sole consequence of Sod1 knockdown in CySC, we manipulated CySC numbers using UAS-Ci and assessed their impact on neighbouring GSCs. Ci overexpression driven by Tj-Gal4 led to a higher number of Tj+ cells (Fig. S4N, and R’’) and a complete loss of Vasa+ GSCs (Fig. S4O, and R’’’). However, the effect on CySC and GSC population was less severe when driven by C-587 Gal4 (Fig.S4N, O, Q’’and Q’’’).

Altered CySC redox state affects niche maintenance signals.

Earlier version: (A-B) The expression of cell polarity marker Dlg (red) in control and Sod1i stem-cell niche (A’’-B’’), (C-D) Representative image showing pErk distribution parallelly with Tj+ cells and its expression pattern through Fire LUT (C’’-D’’). Scale bar – 10 µm. (E) Segregated pErk expression patterns in CySCs near the hub and differentiating cyst cells (CC), mean ± s.e.m, n = 200. (F) Estimation of the differentiation zone through quantification of high pErk expressing regions. (G-I) Activation of PI3K/Tor in somatic lineage was adjudged through reporter p4E-BP in Tj>Sod1i (G’-H’) and levels of the active form was quantified (I), MFI – mean fluorescence intensity shown as mean ± s.e.m, n = 40. (J) Induction of Cyclin D in Sod1i represented as fold change over controls by mean ± s.e.m, n = 3. (K-L) The number of Vasa+ (K) and Tj+ (L) cells across control, Sod1i, Sod1i/PI3KDN (dominant negative allele) and Sod1i/Hhi rescue samples depicted as mean ± s.e.m, n = 10. (M’-N’) Representative image showing distribution of Patched (Ptc) in Tj+ cells in control and Tj>Sod1i. (O-P) Quantification of Ptc (O) and Hh effector Ci (P) expression through fluorescence intensity as mean ± s.e.m, n (Ptc) = 90, n (Ci) = 200. See also Figure S4.

Revised version: (A-B) The expression of cell polarity marker Dlg (red) in control and Sod1i stem-cell niche (A’’-B’’), (C-D) Representative image showing pErk distribution parallelly with Tj+ cells and its expression pattern through Fire LUT (C’’-D’’). Scale bar - 10 µm. Mean florescence intensity (MFI) corresponding to Dlg level (E) was quantified, n = 75, ***P(unpaired t-test)<0.0001. (F) MFI of total pErk expression was quantified, n =200, ***P(unpaired t-test)<0.0001. (G’-H’) Representative image showing distribution of Patched (Ptc) in Tj+ cells in control and Tj>Sod1i. (I-J) Quantification of Ptc (I) and Hh effector Ci (J) expression through fluorescence intensity as mean ± s.e.m, n (Ptc) = 90, n (Ci) = 200, ***P(unpaired t-test)<0.0001 (K-L) The number of Vasa+ (K) and Tj+ (L) cells across control, Sod1i, and Sod1i/Hhi rescue samples depicted as mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001. See also Figure S4.

Elevating CySC antioxidant defence promotes GSC self-renewal

Earlier version: To further substantiate the role of ROS in coordinating the multi-lineage stem populations, we strengthened the cellular defences against oxidants by overexpressing Sod1 in CySCs and monitored the relative GSC/CySC populations. Indeed, we observed an increase in the number of Vasa+ cells (Fig. 5A and S5A’’-B’’), together with a slight reduction in Tj+ cells (Fig. 5B and S5E’’-F’’). Any morphological anomalies associated with cell-cell adhesion or abnormal cellular dispersion, was not observed (Fig. S5A-B and E-F). Enhancing the levels of Sod1 in GSCs promoted the growth of GSC-like cells (Fig. S5C’’, D’’ and I) but did not have any prominent effect on CySCs (Fig. S5G’’, H’’ and J). The increment in Vasa+ GSCs under CySC-induced low redox conditions was also represented by parallel increase in spectrosome number (Fig. 5C-E). The uptick in GSC number due to scavenging of ROS in CySC, might be a result of delayed differentiation as indicated by displacement of Bam+ zone away from hub; thus, confirming non-cell-autonomous role of CySC ROS in maintaining GSC fate (Fig. 5F-H). The data suggest that reducing redox profile of CySCs differentially affected their own self-renewing propensities along with GSC maintenance and optimum redox state of both the stem-cell populations is controlled by CySCs.

Revised version: To further substantiate the role of ROS in coordinating the stem cell populations, we strengthened the cellular defences against oxidants by overexpressing Sod1 in CySCs. We checked the overall ROS level using DHE (Fig. S5I) and monitored the relative GSC/CySC populations. Together with reduction in superoxide levels, we observed an increase in the number of Vasa+ cells (Fig. 5A and S5A’’-B’’), and a slight reduction in Tj+ cells (Fig. 5B and S5E’’-F’’). Any morphological anomalies associated with cell-cell adhesion or abnormal cellular dispersion, was not observed (Fig. S5A-B and E-F). Enhancing the levels of Sod1 in GSCs promoted the growth of GSC-like cells (Fig. S5C’’, D’’ and J) but did not have any prominent effect on CySCs (Fig. S5G’’, H’’ and K). The increment in Vasa+ GSCs under CySC-induced low redox conditions was also represented by a parallel increase in spectrosome number (Fig. 5C-E). The uptick in GSC number due to scavenging of ROS in CySC, might be a result of delayed differentiation as indicated by displacement of Bam+ zone away from hub; thus, confirming non-cell-autonomous role of CySC ROS in maintaining GSC fate (Fig. 5F-H). The data suggest that reducing redox profile of CySCs differentially affected their own self-renewing propensities along with GSC maintenance and optimum redox state of both the stem-cell populations is controlled by CySCs.

Low CySC ROS sustains GSC maintenance

Earlier version: (A-B) Bar graphs illustrating variations in Vasa+ GSC (A) and Tj+ CySCs (B) numbers upon overexpressing Sod1 (Tj>SodOE), represented as mean ± s.e.m, n = 10. (C-E) The spectrosomes labelled with α-spectrin (α-spec) were imaged (C’-D’) and their number was quantified (E). (F-H) The initiation of gonialblast differentiation was determined by measuring the distance of Bam+ zone from hub (red), n = 10 (F), (G’’-H’’) shows distribution of Bam+ reporter expressing cells. Scale bar – 10 µm. See also Figure S5.

Revised version: (A-B) Bar graphs illustrating variations in Vasa+ GSC (A) and Tj+ CySCs (B) numbers upon overexpressing Sod1 (Tj>SodOE), represented as mean ± s.e.m, n = 10, ***P(unpaired t-test)<0.0001, *P(unpaired t-test)<0.01. (C-E) The spectrosomes labelled with α-spectrin (α-spec) were imaged (C’-D’) and their number was quantified (E), n = 10, **P(unpaired t-test)<0.001. (F-H) The initiation of gonialblast differentiation was determined by measuring the distance of Bam+ zone from hub (red), n = 10, *P(unpaired t-test)<0.01 (F), (G’’-H’’) shows distribution of Bam+ reporter expressing cells. Scale bar - 10 µm. See also Figure S5.

Added in the revised version: Proposed role of intercellular redox balance in germline maintenance and differentiation

Added in the revised version: GSC and CySC are associated with different mitochondria abundance which correspond to the presence of a redox differential between two stem cell populations. Disequilibrated redox potential among the two populations due to depletion of Superoxide dismutase, enhances CySC proliferation and precocious GSC differentiation, thereby disrupting the homeostatic balance between the two stem populations.

Discussion

Earlier version: ROS is a crucial mediator of stem cell maintenance where these species act as second messengers to induce different post-translational protein modifications thereby, affecting cell fate45,46. Effectuation of ROS is mainly considered to be autocrine in nature where generation of these oxidative species is compartmentalized with their site of action. Recent studies have reported non-autonomous generation of ROS by NOX enzymes or upon stimulation by growth factors, to play a critical role in regenerative growth47–49. However, in either of the cases, the target cell serves as the source and consequent response. Our observations in multi-lineage Drosophila stem cell niche suggested a previously unaccounted transmission of redox signals inter-cellularly that act as an organizer of niche homoeostasis. The higher ROS threshold of CySCs and resultant CySC-GSC oxidative gradient maintained the physiological redox state of GSC. Reduced ROS in GSC was accompanied with its proliferation50, phenocopied when Sod1 was overexpressed in its neighbours (Fig. 6).

Revised version: ROS is a crucial mediator of stem cell maintenance where these species act as second messengers to induce different post-translational protein modifications thereby, affecting cell fate48,49. The effect of ROS is mainly considered to be autocrine in nature, where generation of these oxidative species is usually restricted to specific cellular compartments, ensuring they act close to their targets. Recent studies have reported non-autonomous generation of ROS by NOX enzymes or upon stimulation by growth factors, to play a critical role in regenerative growth50-52. However, in either of the cases, the target cell serves as the source and the site of response. Among ROS, hydrogen peroxide (H2O2) stands out due to its stability and membrane permeability53, allowing it to diffuse from its source54. Our observations in Drosophila testicular stem cell niche suggested a previously unaccounted role of superoxide dismutases in maintaining niche homoeostasis. The redox perturbations affected cyst stem cell dynamics which indirectly influenced the germline. The higher ROS threshold of CySCs maintained the physiological redox state of GSCs. Reduced ROS in GSC was accompanied with its proliferation55, phenocopied when Sod1 was overexpressed in its neighbours (Fig. 5). However, we feel that net alterations in GSC state were the combinatorial effect of oxidative stress and niche alterations brought about by CySC deregulation.

Removed in the revised version: Proposed role of intercellular redox gradients in GSC maintenance and differentiation

Removed in the revised version: Mitochondrial abundance in GSC-CySC boundary maintains redox differential between two stem cell population. Disequilibrated intercellular redox gradient enhances precocious GSC differentiation and enhances CySC proliferation, thereby disrupting the homeostatic balance between the two stem populations.

Earlier version: The GSCs are attached to the hub via adherens junction to receive the maintenance signals like, Upd that trigger the JAK-STAT pathway27,28, suppression of which activates Bam, acting as a switch from transit amplifying cells to spermatocyte differentiation51,52. E-cadherin levels ensure only undifferentiated stem cells to persist within the niche, while differentiating GSCs lose the competition for niche occupancy25. Elevated ROS in CySC possibly affected STAT phosphorylation through S-glutathionylation53,54, causing reduced presence of dependent transcripts, Socs36E and E-cadherin thereby, resulting in detachment of GSCs from the neighbours as under E-cadherin depleted conditions55–58. This resulted in compromised GSC-CySC paracrine receptions such as EGFR which play a crucial role in maintenance of cell-polarity that segregates self-renewing CySC populations from the ones receiving differentiation cues 39,59,60, leading to accumulation of cells co-expressing both stemness and differentiation markers. In contrast to GSCs, ROS imbalance in CySCs induced accelerated differentiation due to deregulation of PI3K/Tor and Hedgehog pathways. Although, redox modulation of PI3K pathway components has been previously demonstrated61, in this study we found Hedgehog to be probably susceptible to redox regulation.

Revised version: GSCs are anchored to the hub cells via adherens junctions, allowing them to receive crucial maintenance signals28. These include Bone Morphogenetic Proteins (BMPs) such as Decapentaplegic (Dpp) and Glass bottom boat (Gbb), which activate the BMP pathway31 required for GSC self-renewal, and Unpaired (Upd), which triggers the JAK-STAT pathway to promote GSC adhesion to the hub56. The suppression of these pathways activates Bam, acting as a switch from transit-amplifying cells to spermatocyte differentiation57,58. E-cadherin levels ensure only undifferentiated stem cells persist within the niche, while differentiating GSCs lose the competition for niche occupancy26. Elevated ROS in CySC possibly affected STAT phosphorylation through S-glutathionylation59,60, reducing the levels of STAT-dependent transcripts such as Socs36E61, Ptp61F and E-cadherin. This depletion compromises GSC adhesion, resulting in detachment from neighboring cells, similar to E-cadherin loss conditions. 56,62-64. This resulted in compromised GSC-CySC paracrine receptions such as EGFR which play a crucial role in the maintenance of cell-polarity that segregates self-renewing CySC populations from the ones receiving differentiation cues 43,65,66, leading to the accumulation of cells co-expressing both stemness and differentiation markers. In contrast to GSCs, ROS imbalance in CySCs induced accelerated proliferation as well as differentiation due to deregulation of EGFR and Hedgehog pathways. Although, redox modulation of EGFR pathway components has been previously demonstrated67, in this study we found Hedgehog to be probably susceptible to redox regulation. The redox dependent induction of Hedgehog probably, contributed to enhancement in CySC numbers which partly contributed to depletion of GSCs.

Earlier version: However, we do not rule out the possibilities of hub cells playing an equivalent role in GSC maintenance, given their parallel effect in suppressing GSC differentiation28. We had tried using E-132 driver to express Sod1i lines in hub cells but were challenged by substantial adverse effect on fly viability. We also avoided usage of Gal80 dependent clonal populations to maintain a homogenous genetic background and prevent false readouts due to diffusion of ROS signals in unaltered neighbourhood.

Revised version: However, we do not rule out the possibility of hub cells playing an equivalent role in GSC maintenance, given their parallel effect in suppressing GSC differentiation30. Tj-Gal4 depleting Sod1 in both the hub and CySCs, presented a stronger phenotype that CySC specific driver C587-Gal4, indicating potential hub contributions. We had also tried using E-132 Gal4 particularly expressed in the hub to express Sod1i lines in hub cells but were challenged by substantial adverse effect on fly viability. While ROS accumulation typically enhances EGFR signaling promoting premature GSC differentiation68, the loss of EGFR in somatic cells increases the number of GSCs65 and reduced EGFR activity leads to fewer somatic cells 69. However, we feel that the response of germline to EGFR modulation is more context-dependent as Sod1 depletion caused parallel reduction of pErk level and GSC number. We tried an EGFR gain-of-function rescue of CySC depletion but the driven progenies were lethal in the absence of Sod1. We also avoided usage of Gal80 dependent clonal populations to maintain a homogenous genetic background and prevent false readouts due to diffusion of ROS signals in unaltered neighbourhood.

Earlier version: The proposed concept of intercellular ROS communication can be of importance in deciphering the biochemical adaptability and plasticity of different niches influencing stem cell fate, ensuring niche size and architecture to prevent stem cell loss and aging. This has been observed in neural stem cells where inflammation-induced quiescence recovers the regenerating capacity of aging brain62. In addition to metabolic variations, dependence of the germline on somatic neighbours for its redox state might be one the reasons behind its presumed immortal nature and renitence to aging63. The same can also be applied to cancer stem cells which maintain niche occupancy and resist oxidative stress induced apoptosis probably, by receiving redox cues from the environment. However, further work is required to elucidate the myriad of driver mechanisms intersecting at the realm of redox regulation that extend beyond the present system into broader translational areas.

Revised version: The proposed concept of intercellular ROS communication can be of importance in deciphering the biochemical adaptability and plasticity of different niches influencing stem cell fate, ensuring niche size and architecture to prevent stem cell loss and aging. This has been observed in neural stem cells where inflammation-induced quiescence recovers the regenerating capacity of aging brain70. In addition to metabolic variations, dependence of the germline on somatic neighbours for its redox state might be one the reasons behind its presumed immortal nature and resistance to aging71. The same can also be applied to cancer stem cells, which maintain niche occupancy and resist oxidative stress-induced apoptosis, probably by receiving redox cues from the environment. However, further work is required to elucidate the myriads of driver mechanisms intersecting in the realm of redox regulation that extend beyond the present system into broader translational areas.

Earlier version: Materials and Methods Fly strains

Revised version: Materials and methods

Added in the revised version: Fly strains

Earlier version: Fly stocks were maintained and crosses were set at 25 °C on normal corn meal and yeast medium unless otherwise indicated. All fly stocks (BL-24493) UAS SOD1 RNAi, (BL-32909) UAS SOD1 RNAi, (BL-29389) UAS SOD1 RNAi, (BL-32983) UAS SOD2 RNAi, (BL-24754) UAS SOD1, (BL-32489) UAS hhRNAi, (BL-8288) UAS PI3KDN, (BL-25751) UAS Dcr-nos, (BL-55122) UAS-FUCCI, were obtained from Bloomington Drosophila Stock Centre. Nos-Gal4, Tj-Gal4, Tj-Gal4-BamGFP, gstD1-GFP, TFAM-GFP were kind gift from U. Nongthomba, P. Majumder, K. Ray, B. C. Mandal, Hong Xu labs respectively. For optimal Gal4 activity parent crosses were shifted to 29 °C till the eclosion of adults. 3-5 days old males were taken for all the experiments. The control lines for all experiments are the corresponding Gal4-line crossed with wild type OregonR+.

Revised version: Fly stocks were maintained and crosses were set at 25°C on normal corn meal and yeast medium unless otherwise indicated. All fly stocks (BL-24493) UAS SOD1 RNAi, (BL-32909) UAS SOD1 RNAi, (BL-29389) UAS SOD1 RNAi, (BL-32983) UAS SOD2 RNAi, (BL-24754) UAS SOD1, (BL-32489) UAS hhRNAi, (BL-25751) UAS Dcr-nos Gal4, (BL-55122) UAS-FUCCI, were obtained from Bloomington Drosophila Stock Centre. Nos-Gal4, Tj-Gal4, Tj-Gal4-BamGFP, gstD1-GFP, TFAM-GFP, UAS-Ci were kind gift from U. Nongthomba, P. Majumder, K. Ray, B. C. Mandal, Hong Xu, LS Shahidhara labs respectively. Parents were maintained at 25°C, and crosses were set at the same temperature. To ensure optimal Gal4 activity, progenies were shifted to 29°C until eclosion, with aging also conducted at 29°C. Males aged 3–5 days were used for all experiments. The control lines for all experiments are the corresponding Gal4-line crossed with wild type OregonR+.

Dissection and immunostaining

Earlier version: Anesthetized flies were dissected in 1X Phosphate Buffer Saline (PBS). All incubations were carried out at room temperature (25 °C) unless otherwise mentioned. Testes were fixed in 4% paraformaldehyde for 30 minutes, followed by multiple washes with 0.1% PBTX (PBS + TritonX 100). Post-blocking in 0.5% Bovine Serum Albumin, testis was incubated overnight at 4°C for primary antisera, followed by washing in 0.1% PBTX 3 times 15 minutes each before incubating with secondary antibody. The tissues were counterstained with DAPI (4,6-diamidino-2-phenylindole) for 20 minutes, followed by three washes in 0.1% PBTX 10 minutes each. For PTC staining, a modified protocol utilizing PIPES-EGTA buffer was used as described64. The dpErk was labelled by dissecting fly testes in Schneider’s media, followed by fixation in 4% PFA for 30 min, 3x wash in 0.1% PBTX, blocked and incubated in primary antibody overnight; each step supplemented with phosphatase inhibitor cocktail 2 (1:100, Sigma, cat#P5726). Samples were mounted in DABCO anti-fade medium prior imaging.

Revised version: Anesthetized flies were dissected in 1X Phosphate Buffer Saline (PBS). All incubations were carried out at room temperature (25°C) unless otherwise mentioned. Testes were fixed in 4% paraformaldehyde for 30 minutes, followed by three washes with 0.3% PBTX (PBS + TritonX 100). Post-blocking in 0.5% Bovine Serum Albumin, testis was incubated overnight at 4°C for primary antisera, followed by washing in 0.3% PBTX 3 times 15 minutes each before incubating with secondary antibody. The tissues were counterstained with DAPI (4,6-diamidino-2-phenylindole) for 20 minutes, followed by three washes in 0.3% PBTX for 15 minutes each. For Patched antibody staining, a modified protocol utilizing PIPES-EGTA buffer was used as described72. For Eya staining, the tissues were incubated in primary antibody for 2 days. The dpErk was labelled by dissecting fly testes in Schneider’s media, followed by fixation in 4% PFA for 30 min, 3x wash in 0.1% PBTX, blocked and incubated in primary antibody overnight; each step supplemented with phosphatase inhibitor cocktail 2 (1:100, Sigma, cat#P5726). Samples were mounted in DABCO anti-fade medium prior imaging.

Earlier version: The following primary antibodies were used in the experiments – anti-FasIII (7G10) (1:120, DSHB (Developmental Studies Hybridoma Bank)), anti-Eya (1:50, DSHB), anti-DEcad (DCAD2) (1:50, DSHB), anti-α-Spectrin (3A9) (1:20, DSHB), anti-Ptc (1:100,DSHB), anti-Ci (1:50, DSHB), anti-βtubulin (1:300, DSHB), anti-pERK (1:100, Cell Signaling (4370)), ATP5A (1:700, Abcam (ab14748)), anti-Tj (1:5000), anti-Vasa (1:4000), anti-Zfh1(1:2000). The primary labelling was detected using appropriate Alexa-Flour tagged secondary antibody (ThermoFisher Scientific).

Revised version: The following primary antibodies were used in the experiments - anti-FasIII (7G10) (1:120, DSHB (Developmental Studies Hybridoma Bank)), anti-Eya (1:50, DSHB), anti-DEcad (DCAD2) (1:50, DSHB), anti-α-Spectrin (3A9) (1:20, DSHB), anti-Ptc (1:50,DSHB), anti-Ci (1:50, DSHB), anti -Dlg (1:50), anti-βtubulin (1:300, DSHB), anti-pERK (1:100, Cell Signaling (4370)), ATP5A (1:700, Abcam (ab14748)), anti-Tj (1:5000), anti-Vasa (1:4000), anti-Zfh1(1:2000). The primary labelling was detected using appropriate Alexa-Flour tagged secondary antibody (ThermoFisher Scientific).

Earlier version: To assess superoxide levels, testes were dissected in Schneider’s medium (SM) and immediately incubated in 30 µM Dihydroxyethidium (DHE) at 25 °C in dark. Testes were washed thrice with SM for 7 minutes each before quantifying the emitted fluorescence using SpectraMax iD5 (Molecular Devices).

Revised version: To assess superoxide levels, testes were dissected in Schneider’s medium (SM) and immediately incubated in 30 µM Dihydroxyethidium (DHE) at 25 °C in the dark. Testes were washed thrice with SM for 7 minutes each before quantifying the emitted fluorescence using SpectraMax iD5 (Molecular Devices).

Quantitative Reverse transcription–PCR

For semi-quantitative RT-PCR and qRT-PCR analyses, total RNA was isolated from testes of 3-5 days old male flies using Qiagen RNA extraction kit. RNA pellets were resuspended in nuclease-free water and quantity of RNA was spectrophotometrically estimated. First-strand cDNA was synthesized from 1-2 µg of total RNA as described earlier. The prepared cDNAs were subjected to real time PCR using forward and reverse primer pairs as listed below, using 5 µl qPCR Master Mix (SYBR Green, ThermoFisher Scientific), 2 pmol/µl of each primer per reaction for 10 µl final volume in ABI 7500 Real time PCR machine. The fold change in expression was calculated through 2-ΔΔCt method. The primers used are listed in Supplementary Table S1.

Imaging and image analysis

Earlier version: Confocal imaging was carried out using Zeiss LSM 900 confocal microscope, using appropriate dichroics and filters. Images were further analyzed and processed for brightness and contrast adjustments using ImageJ (Fiji). Mean intensity measurements were carried out using standard Fiji plug-ins, 2D surface plot quantification or Zeiss Zen software as indicated. Quantification of cell number/nuclei in single or co-stained populations (Tj, Vasa, Zfh1, Eya, FUCCI) were performed from individual slices of z-stacks using semi-automated Cell Counter plugin in Fiji for all the mentioned genotypes. The relative distance from the two reference points in the images was estimated through Inter-edge Distance ImageJ Macro v2.0 from Github. All images were assembled using Adobe Photoshop version 21.2.1.

Revised version: Confocal imaging was carried out using Zeiss LSM 900 confocal microscope, using appropriate dichroics and filters. Images were further analyzed and processed for brightness and contrast adjustments using ImageJ (Fiji). Mean intensity measurements were carried out using standard Fiji plug-ins,The relative distance from the two reference points in the images was estimated through Inter-edge Distance ImageJ Macro v2.0 from Github. All images were assembled using Adobe Photoshop version 21.2.1.

Added in the revised version: Quantification of GSCs, CySCs and CCs

Added in the revised version: For GSC quantification, only cells in direct contact with the hub were considered. While single slices are shown for representation, counting was performed on individual slices of z-stacks using the semi-automated Cell Counter plugin in Fiji for all specified genotypes.

Added in the revised version: For CySC quantification, all Tj-positive and Zfh1-positive cells present in each testis were counted and represented. For cyst cells quantification, only those near the hub were considered, which does not reflect the total number of Eya-positive cells. All quantifications were performed using the Cell Counter plugin in Fiji.

Added in the revised version: Mitochondrial and GstD-GFP quantification

Added in the revised version: The mitochondrial analyzer first generates 2D mitochondrial regions of interest (ROIs), followed by the measurement of their total area and circularity. Initially, this is done using a thresholded image of a single slice. For GstD-GFP quantification, we have manually demarcated the cell, taking Dlg as a reference, and quantified the GFP intensity.

Immunoblot Analysis

Protein was extracted from the dissected testes using RIPA buffer as described previously and quantified using Bradford reagent (Biorad). Equivalent concentration of lysate was denatured using 1X Laemelli Buffer with 1 M DTT at 95°C, separated in 10% SDS-PAGE and transferred onto PVDF membranes. Membranes were blocked with commercial blocking solution in TBST base before sequential primary antibody incubation with anti-Vasa and anti-β-Tubulin (DSHB, E7). Secondary detection was performed using HRP-tag anti-mouse antibody (GE Amersham).

Statistical analyses

Earlier version: All data points represent at least three biological replicates and the sample sizes carrying adequate statistical power are mentioned in the figure legends. Statistical significance for each experiment were calculated using two tailed Student’s t-test unless otherwise mentioned, using GraphPad Prism 8 software. The P values were calculated through pairwise comparison of the data with wild type or driver alone and driven RNAi lines. Significance values for sample sizes mentioned in figure legends were represented as *P < 0.01, **P <0.001, or ***P < 0.0001.

Revised version: The sample sizes carrying adequate statistical power are mentioned in the figure legends. Statistical significance for each experiment was calculated using two-tailed Student’s t-test, unless otherwise mentioned, using GraphPad Prism 8 software. The P values were calculated through pairwise comparison of the data with wild type or driver alone and driven RNAi lines. Significance values for sample sizes mentioned in figure legends were represented as *P < 0.01, **P <0.001, or ***P < 0.0001.

Removed in the revised version: Author contribution

Removed in the revised version: D.S. and O.M conceived the study and designed the experiments. O.M, A.C and T.C performed the experiments. D.S and O.M analyzed the data. D.S, O.M. and A.C. wrote the manuscript.

Removed in the revised version: Supporting information